<quiz>
  <question type="coderunner">
    <name>
      <text>Name</text>
    </name>
    <questiontext format="html">
      <text>
<![CDATA[TESTE]]>
</text>
    </questiontext>
    <generalfeedback format="html">
      <text />
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber />
    <coderunnertype>cpp_program</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0, 0, 10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload />
    <useace />
    <resultcolumns />
    <template />
    <iscombinatortemplate />
    <allowmultiplestdins />
    <answer>Solution
<![CDATA[ola eu sou uma solu&#231;&#227;o]]>
</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre />
    <language />
    <acelang />
    <sandbox />
    <grader />
    <cputimelimitsecs />
    <memlimitmb />
    <sandboxparams />
    <templateparams />
    <hoisttemplateparams>1</hoisttemplateparams>
    <twigall>0</twigall>
    <uiplugin />
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex />
    <filenamesexplain />
    <displayfeedback>1</displayfeedback>
    <testcases>
    </testcases>
  </question>

</quiz>