<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 933  -->
  <question type="coderunner">
    <name>
      <text>Fizz Busão</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>
    Jonnie Ruquer gastava muito tempo no translado entre sua casa e a universidade, e decidiu criar um novo passatempo baseado em um de seus jogos favoritos. A cada veículo que passava, ele contava a frequência do tipo e gritava um descritor. Para toda terceira
    ocorrência de qualquer veículo, soltava "fizz" e a cada quinta ocorrência de um ônibus, era "busao". A diversão era plena quando estas condições coincidiam, e ele urrava "fizzbusao". Nos demais casos, ele simplesmente dizia a quantidade de veículos
    e torcia pelo próximo grito. Para verificar se ele está seguindo as regras, crie um programa que, dada uma sequência de veículos, imprima a sequência de descritores conforme as regras dadas.
</p>
<p>
    <b>Entrada</b><br> A entrada consiste de uma linha contendo <i>N</i> (<i>1 ≤ N ≤ 1000</i>) caracteres contíguos, no alfabeto <tt>{C, O}</tt>, indicando a ocorrência de um carro ou de um ônibus, respectivamente.
</p>
<p>
    <b>Saída</b><br> A saída deve ser composta por <i>N</i>linhas, cada uma apresentando o descritor da situação de cada veículo, na mesma ordem em que são dados.
</p>
<p>
    <b>Notas</b><br> 
  No primeiro exemplo, como não há ônibus apenas a condição de "fizz" é atendida.
  No segundo exemplo, há veículos suficientes para o "fizz", mas não para o "busao".
  No último exemplo, o décimo ônibus possibilita a o grito "busao", e a passagem do décimo quinto veículo permite que Jonnie finalmente solte seu urro.
  
</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>c_program</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0, 0, 10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[#include <stdio.h>

int main() {
    int veiculos = 0, busao = 0;
    char c = 'O';
    
    while (1) {
        scanf("%c", &c);
        
        if (c != 'O' && 'C' != c)
            break;
        
        ++veiculos;
        if (c == 'O')
            ++busao;
            
        if (0 == (veiculos % 3)) {
            if (c == 'O' && 0 == (busao % 5)) {
                printf("fizzbusao\n");
            } else {
                printf("fizz\n");
            }
        } else {
            if (c == 'O' && 0 == (busao % 5)) {
                printf("%d\n", busao);
            } else {
                printf("%d\n", veiculos);
            }
        }
    }

    return 0;
}]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs>1</cputimelimitsecs>
    <memlimitmb>256</memlimitmb>
    <sandboxparams><![CDATA[{"linkargs":["-lm"]}]]></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>CCCCC
</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
5
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>COCO</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>OOOOOCCCCCOOOOO
</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
5
fizz
7
8
fizz
10
11
fizz
13
14
fizzbusao
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>C
</text>
      </stdin>
      <expected>
                <text>1
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>O
</text>
      </stdin>
      <expected>
                <text>1
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>CCCCOCCCOOOOOOCCCOCCOOCOCOCOCCOOOCCCOCOOCOOCOOCCOOOOCCOCOCOOCCCCCCOCCOOOCCOOCOOOOOOCCCOOCOOOCOOCCCOCCOCOOOOOCCCCCOCCOOCOCCOOCOOOCOCCCOOOCCCCCCCOCCOOOOCCOOCCOOOOCCOCCCOOCCOCCCOCCOOCCCCCCCOOCCOOOCOOOOCOCOOCCOCOCOOCCOOOCCCCCCOOCOOCOCOOCCCOCCCCOCOCCOCOCCCCCCOOCCCOOOCCOOOOCCCCCCCCOOCCCCCCOOCOCCCCOCCCCOCCCCOOCCCOOCCOCCOOCCOOOOOOOCOCCCOCOOCCCCCOOCOCOOOOOCCCCCOOOCCCCOCOCCOCOOOCOCOOCCOCOCOOCCCOOCOCCCCCCOCOOOOOOCOCOCOOCOOOCOOCCOCOOCOCCCCOCOCCCCOOCCOOOCOOCOCCCCOOOCOOOOOOOOOOOOCOOCOOCOOOOCOCOOCOCOOCCCOOCOCCCOCCCCCOOOCOCOCCOCCOCOOOOOCOOCOOOCOCCOOCCOOOCCCOOCOOOOOCCOOCOOOOOOOOCOOOOOOCCCOCOCCOOOOCCOOCCOOOOOCCOOCOCOCOOOCCCCOOCCCOOOCOOOOCOOOCCOCOCOOOCOCCCOCCOOOCCOCCOCOCCOCOCOOOCOCCCCCOOOCOCOOOCCCOOOOCCOOCOOOOOOOOCCCCOCCCCCCCCOCOCOCCOOOOCOOOOOOOOCCCCOOOCOOCOCCCCOOCOCCCOCCCCOOCCCCOOOCCCOCOCOCCCOOOCCCCOCOCOOOOOOCCOOOOCCOCCCCCCCCCOCCCOOOOCCCCOOOOOOOOCOCCOCCCOOOOOCOCCCCOOCCOCOCCOOOOCOCCCOOOOCCCOCCCOCCOCOCCCOOOOOCOCCCCCCOCOCOOOOOOOOCCCCCCOCOCOCOOCCCOOOOCOOCOOOOOOOCCOOOCOCOCCOOCOOCOCCCOCOCOCOOCCOCCOOCOOOOOOOOC
</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
5
fizz
7
8
fizz
10
11
fizzbusao
13
14
fizz
16
17
fizz
19
20
fizz
10
23
fizz
25
26
fizz
28
29
fizz
31
15
fizz
34
35
fizz
37
38
fizz
40
41
fizzbusao
43
44
fizz
46
47
fizz
49
25
fizz
52
53
fizz
55
56
fizz
58
30
fizz
61
62
fizz
64
65
fizz
67
68
fizz
70
71
fizzbusao
73
74
fizz
76
77
fizz
79
40
fizz
82
83
fizz
85
86
fizz
45
89
fizz
91
92
fizz
94
50
fizz
97
98
fizz
100
101
fizz
103
104
fizz
55
107
fizz
109
110
fizz
112
113
fizz
115
116
fizz
60
119
fizz
121
122
fizz
124
125
fizz
65
128
fizz
130
131
fizz
133
134
fizz
70
137
fizz
139
140
fizz
142
143
fizz
145
146
fizz
148
149
fizzbusao
151
152
fizz
154
155
fizz
157
158
fizzbusao
160
161
fizz
163
164
fizz
166
167
fizz
169
170
fizzbusao
172
173
fizz
175
176
fizz
178
179
fizz
181
182
fizz
184
185
fizz
187
90
fizz
190
191
fizz
193
194
fizz
95
197
fizz
199
200
fizz
202
100
fizz
205
206
fizz
208
209
fizz
211
212
fizz
105
215
fizz
217
218
fizz
220
221
fizz
223
224
fizz
110
227
fizz
229
230
fizz
232
233
fizz
235
115
fizz
238
239
fizz
241
242
fizz
244
245
fizz
247
248
fizz
250
251
fizz
253
254
fizzbusao
256
257
fizz
259
260
fizz
262
263
fizz
125
266
fizz
268
269
fizz
271
272
fizz
274
275
fizz
277
130
fizz
280
281
fizz
283
284
fizz
286
287
fizz
289
290
fizz
292
293
fizz
295
296
fizz
135
299
fizz
301
302
fizz
304
305
fizz
307
308
fizz
310
311
fizzbusao
313
314
fizz
316
317
fizz
319
320
fizzbusao
322
323
fizz
325
326
fizzbusao
328
329
fizz
331
332
fizz
334
335
fizz
337
338
fizz
340
155
fizz
343
344
fizz
346
347
fizzbusao
349
350
fizz
352
353
fizz
355
356
fizz
358
359
fizz
361
165
fizz
364
365
fizz
367
368
fizz
370
170
fizz
373
374
fizz
376
377
fizz
379
380
fizzbusao
382
383
fizz
385
386
fizz
388
389
fizz
180
392
fizz
394
395
fizz
397
398
fizz
400
401
fizz
185
404
fizz
406
407
fizz
409
410
fizzbusao
412
413
fizz
415
416
fizz
195
419
fizz
421
422
fizz
424
425
fizz
200
428
fizz
430
431
fizz
433
434
fizz
436
437
fizz
439
440
fizz
442
205
fizz
445
446
fizz
448
449
fizzbusao
451
452
fizz
454
455
fizz
457
458
fizz
215
461
fizz
463
464
fizzbusao
466
467
fizz
469
225
fizz
472
473
fizz
475
476
fizz
230
479
fizz
481
482
fizz
484
235
fizz
487
488
fizz
490
491
fizz
493
494
fizzbusao
496
497
fizz
499
500
fizz
502
503
fizz
505
506
fizz
508
245
fizz
511
512
fizz
514
515
fizz
517
518
fizz
250
521
fizz
523
524
fizz
255
527
fizz
529
530
fizz
532
260
fizz
535
536
fizz
538
539
fizz
541
542
fizzbusao
544
545
fizz
547
548
fizz
550
551
fizzbusao
553
554
fizz
556
557
fizz
275
560
fizz
562
563
fizz
280
566
fizz
568
569
fizz
285
572
fizz
574
575
fizz
577
578
fizzbusao
580
581
fizz
583
584
fizz
586
295
fizz
589
590
fizz
592
593
fizz
595
300
fizz
598
599
fizz
601
602
fizz
305
605
fizz
607
608
fizz
610
611
fizz
613
614
fizzbusao
616
617
fizz
619
620
fizz
622
623
fizzbusao
625
626
fizz
628
629
fizzbusao
631
632
fizz
634
635
fizz
637
638
fizzbusao
640
641
fizz
643
644
fizz
646
647
fizz
649
330
fizz
652
653
fizz
655
656
fizz
658
659
fizz
661
335
fizz
664
665
fizz
667
668
fizz
340
671
fizz
673
674
fizz
676
677
fizz
679
680
fizz
345
683
fizz
685
686
fizz
688
689
fizzbusao
691
692
fizz
694
695
fizz
697
355
fizz
700
701
fizz
360
704
fizz
706
707
fizz
709
710
fizz
712
713
fizz
715
716
fizz
718
719
fizz
721
365
fizz
724
725
fizz
727
728
fizz
370
731
fizz
733
734
fizzbusao
736
737
fizz
739
740
fizz
742
743
fizzbusao
745
746
fizz
748
749
fizz
751
752
fizz
754
385
fizz
757
758
fizz
760
761
fizz
763
764
fizz
766
767
fizz
769
770
fizz
390
773
fizz
775
776
fizz
778
779
fizz
781
395
fizz
784
785
fizz
787
788
fizz
790
791
fizz
793
794
fizzbusao
796
797
fizz
799
800
fizzbusao
802
803
fizz
805
806
fizz
410
809
fizz
811
812
fizz
814
815
fizz
817
818
fizz
820
821
fizz
823
824
fizz
826
415
fizz
829
830
fizz
832
833
fizz
835
420
fizz
838
839
fizz
841
425
fizz
844
845
fizz
847
848
fizz
850
851
fizzbusao
853
854
fizz
856
857
fizz
859
860
fizz
862
863
fizzbusao
865
866
fizz
868
869
fizz
871
440
fizz
874
875
fizz
877
878
fizz
880
445
fizz
883
884
fizz
886
887
fizz
889
890
fizz
892
893
fizz
895
896
fizz
450
899
fizz
901
902
fizz
455
905
fizz
907
908
fizz
910
911
fizz
913
914
fizz
916
460
fizz
919
920
fizz
465
923
fizz
925
926
fizz
928
929
fizz
931
932
fizz
934
935
fizzbusao
937
938
fizz
940
941
fizz
943
944
fizzbusao
946
947
fizz
949
950
fizzbusao
952
953
fizz
955
956
fizz
485
959
fizz
961
962
fizz
964
965
fizz
490
968
fizz
970
971
fizz
973
974
fizz
976
977
fizzbusao
979
980
fizz
982
983
fizz
985
986
fizz
988
500
fizz
991
992
fizz
994
505
fizz
997
998
fizz
1000
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>OOCCCCCOCOOOCCCOOOOCCCCCCCOOOOOCCCCOCOCCCCCOOCOCOOOCOOOCCOOCCOOOCCOCCCOCOOCCCCOOCCOCOCCCOCOCOCCCOCOOCOCCCCOCOOOOOCCOCCCCOCCCCOOOCOCCOOOOCCCOCOCOOCOCCOOCCOOOCOCOCOCCOCOCCCOCOOOCOCOCCCOCCOOCCOOCCOCOOCOCCOCCOCOOOCOCOCOCCCOCCOCCCCOCCCCCCCCC</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
5
fizz
7
8
fizz
10
5
fizz
13
14
fizz
16
17
fizz
10
20
fizz
22
23
fizz
25
26
fizz
28
29
fizz
15
32
fizz
34
35
fizz
37
38
fizz
40
41
fizz
43
44
fizz
46
20
fizz
49
50
fizz
52
53
fizzbusao
55
56
fizz
58
59
fizz
61
62
fizzbusao
64
65
fizz
67
68
fizz
70
71
fizz
73
35
fizz
76
77
fizz
79
80
fizz
82
83
fizz
85
86
fizz
88
40
fizz
91
92
fizz
94
95
fizz
97
98
fizz
45
101
fizz
103
104
fizz
106
107
fizz
109
110
fizzbusao
112
113
fizz
115
116
fizz
118
119
fizz
121
122
fizz
124
125
fizzbusao
127
128
fizz
130
131
fizz
133
60
fizz
136
137
fizz
139
140
fizz
142
143
fizzbusao
145
146
fizz
148
149
fizz
151
152
fizz
70
155
fizz
157
158
fizz
160
161
fizzbusao
163
164
fizz
166
167
fizz
169
170
fizz
172
173
fizzbusao
175
176
fizz
178
179
fizz
181
182
fizz
184
185
fizzbusao
187
188
fizz
190
191
fizz
193
194
fizz
90
197
fizz
199
200
fizz
202
203
fizz
205
206
fizzbusao
208
209
fizz
211
212
fizz
214
100
fizz
217
218
fizz
220
221
fizz
223
224
fizz
226
227
fizz
229
230
fizz
232
233
fizz
235
236
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>OOOOOOOOOOOOOOO
</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
5
fizz
7
8
fizz
10
11
fizz
13
14
fizzbusao
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>CCCCCCCCCCCCCCC
</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
5
fizz
7
8
fizz
10
11
fizz
13
14
fizz
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>COOCOCCOOCOCCOCCOOOCOCOOOCOOOOOCCOCOCOOOOCOOOCOOOOCOOCOCCOOOOCCOCCCOCCCCOCOOCOOCOOCCCOCOCOCCCOOCCCCOCOCCCOOCOOCOOOCOCCCOOCCCOOOCOOCCOOCCOOCOCOCOCCOOCCOOOCCOOCOOOCOCCOCCOOCOCCOOCCOCCOCCCCCOOOCCOCCOCCOOCOCOOCOOOOOCCOCCCOCCOCOCCOCOCCCOOOOOCCOCCCCOOOCOOOCCCCCOCCCCCCCCOOCCCOCCCOOOCCOOOCOOOCOOCOOCCCOCCCOCOOOCOOCOCCCOOCCCCOOCCOCCOCOOOOCOOOCCOOOCOOOCCOOCOCCOCCOOCCOCCOCCOOCCCCOCCCCCCCOCOOCCOCOCCCOCCCOOCOOOCOOCCCOCCOOCOOCCCOCCOOOOOOCOOOCOOOOCCCOCOCCOCOOCCCCOOOOOOCOCCOOCOOOOCCOOCOCCCCCCCOOCOOOOOCOCCCCOOCCCCOOCOCCCCOOOCCCCOOOOOOOOOOCCOOCCCOOOCCOCCOCCOOOCOCOOCCCOOCCCOOOOOOOOOOOOCOCCCOOCOOCOOCCOCCCOOOCOCCOCCCCOCOCOOCCCCOOOCCCCCOCCCOOCOCCCCCOOOOOCCOOCOCOCCCCOOOOOCOCCCCCOOOOOOOOOOCCOCOOOCCCOOCOCCOCCCCOOOCOOCOOOCOCOOOCOCCOOCOOOCCOOOOCCCCOCOOOOCOOCCCCOC
</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
5
fizz
7
8
fizzbusao
10
11
fizz
13
14
fizz
16
17
fizz
10
20
fizz
22
23
fizz
25
26
fizzbusao
28
29
fizz
31
32
fizz
20
35
fizz
37
38
fizz
40
25
fizz
43
44
fizz
46
47
fizzbusao
49
50
fizz
52
53
fizz
35
56
fizz
58
59
fizz
61
62
fizz
40
65
fizz
67
68
fizz
70
71
fizz
73
74
fizz
76
77
fizzbusao
79
80
fizz
82
83
fizz
85
86
fizz
50
89
fizz
91
92
fizz
94
95
fizz
97
98
fizz
100
101
fizzbusao
103
104
fizz
106
107
fizz
109
110
fizz
60
113
fizz
115
116
fizz
118
119
fizz
65
122
fizz
124
125
fizz
127
128
fizz
70
131
fizz
133
134
fizz
136
137
fizz
139
75
fizz
142
143
fizz
145
146
fizz
148
149
fizz
80
152
fizz
154
155
fizz
157
158
fizzbusao
160
161
fizz
163
164
fizz
166
167
fizz
90
170
fizz
172
173
fizz
175
176
fizz
178
95
fizz
181
182
fizz
184
185
fizz
187
188
fizz
190
191
fizz
100
194
fizz
196
197
fizz
199
200
fizz
202
203
fizzbusao
205
206
fizz
208
209
fizzbusao
211
212
fizz
214
215
fizz
217
218
fizz
220
221
fizz
115
224
fizz
226
227
fizz
229
230
fizz
232
233
fizzbusao
235
236
fizz
238
239
fizz
241
242
fizz
244
125
fizz
247
248
fizz
250
251
fizz
253
254
fizz
130
257
fizz
259
260
fizz
262
263
fizz
265
266
fizz
268
269
fizz
271
272
fizz
274
135
fizz
277
278
fizz
280
281
fizz
140
284
fizz
286
287
fizz
289
145
fizz
292
293
fizz
295
296
fizz
298
299
fizz
301
150
fizz
304
305
fizz
307
308
fizz
310
311
fizzbusao
313
314
fizz
316
317
fizz
319
320
fizz
322
323
fizz
160
326
fizz
328
329
fizz
331
165
fizz
334
335
fizz
337
338
fizzbusao
340
341
fizz
343
344
fizz
346
175
fizz
349
350
fizz
352
353
fizz
355
356
fizz
358
180
fizz
361
362
fizz
364
365
fizz
367
368
fizz
370
371
fizz
373
374
fizz
376
377
fizz
185
380
fizz
382
383
fizz
385
386
fizz
388
389
fizz
190
392
fizz
394
395
fizz
397
398
fizz
195
401
fizz
403
404
fizz
406
407
fizz
409
410
fizzbusao
412
413
fizz
415
416
fizz
418
419
fizz
421
205
fizz
424
425
fizz
427
210
fizz
430
431
fizz
433
215
fizz
436
437
fizz
439
440
fizz
442
443
fizz
445
220
fizz
448
449
fizz
451
452
fizz
454
225
fizz
457
458
fizz
460
461
fizz
230
464
fizz
466
467
fizz
469
470
fizzbusao
472
473
fizz
475
476
fizz
478
479
fizz
481
482
fizz
484
240
fizz
487
488
fizz
490
245
fizz
493
494
fizz
496
497
fizz
499
500
fizz
502
503
fizz
250
506
fizz
508
509
fizz
511
512
fizz
514
515
fizz
517
255
fizz
520
521
fizz
260
524
fizz
526
527
fizz
529
265
fizz
532
533
fizz
535
536
fizz
538
539
fizz
541
270
fizz
544
545
fizz
547
548
fizz
550
275
fizz
553
554
fizz
556
557
fizz
559
560
fizz
280
563
fizz
565
566
fizzbusao
568
569
fizz
571
290
fizz
574
575
fizz
577
578
fizz
580
581
fizzbusao
583
584
fizz
586
587
fizz
589
590
fizz
592
300
fizz
595
596
fizz
598
599
fizz
601
602
fizz
604
605
fizzbusao
607
608
fizz
610
611
fizz
613
614
fizz
310
617
fizz
619
620
fizz
622
623
fizz
625
626
fizz
628
629
fizz
631
632
fizz
634
315
fizz
637
638
fizz
640
641
fizzbusao
643
644
fizz
646
647
fizz
649
650
fizz
652
325
fizz
655
656
fizz
658
659
fizz
661
662
fizz
330
665
fizz
667
668
fizzbusao
670
671
fizz
673
674
fizz
340
677
fizz
679
680
fizz
682
683
fizz
345
686
fizz
688
689
fizz
691
692
fizz
694
695
fizz
350
698
fizz
700
701
fizz
703
355
fizz
706
707
fizz
709
710
fizz
360
713
fizz
715
716
fizz
718
719
fizzbusao
721
722
fizz
724
725
fizz
727
728
fizz
730
370
fizz
733
734
fizz
736
737
fizzbusao
739
740
fizz
742
743
fizz
745
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
    <tags>
      <tag><text>Difícil</text>
</tag>
      <tag><text>Repetição</text>
</tag>
      <tag><text>C</text>
</tag>
    </tags>
  </question>

</quiz>