<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 1501  -->
  <question type="coderunner">
    <name>
      <text>Fizz Busão</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>
<b></b>
Jonnie Ruquer gastava muito tempo no translado entre sua casa e a universidade, e decidiu criar um novo passatempo baseado em um de seus jogos favoritos. A cada veículo que passava, ele contava a frequência do tipo e gritava um descritor. Para toda terceira ocorrência de qualquer veículo, soltava "fizz" e a cada quinta ocorrência de um ônibus, era "busao". A diversão era plena quando estas condições coincidiam, e ele urrava "fizzbusao". Nos demais casos, ele simplesmente dizia a quantidade de veículos e torcia pelo próximo grito. 

Para verificar se ele está seguindo as regras, crie um programa que, dada uma sequência de veículos, imprima a sequência de descritores conforme as regras dadas.

</p>

<p>
<b>Entrada</b><br />
A entrada consiste de uma linha contendo \( N \) (\( 1 \leq N \leq 1000 \)) caracteres contíguos, no alfabeto \( \{C, O\} \), indicando a ocorrência de um carro ou de um ônibus, respectivamente.

</p>

<p>
<b>Saida</b><br />
A saída deve ser composta por \( N \) linhas, cada uma apresentando o descritor da situação de cada veículo na mesma ordem em que são dados.

</p>

<p>
<b>Notas</b><br />
No primeiro caso, como não há ônibus apenas a condição de "fizz" é atendida.

No segundo exemplo, há veículos suficientes para o "fizz", mas não para o "busao".

No último exemplo, o décimo ônibus possibilita a o grito "busao", e a passagem do décimo quinto veículo permite que Jonnie finalmente solte seu urro.

</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[A solução envolve dois contadores, um para veículos e outro para ônibus, que são incrementados conforme a entrada. Toda vez que o contador de ônibus é múltiplo de 5, Jonnie Ruquer pode dizer "busao". Toda vez que o contador de veículos é múltiplo de 3, ele pode dizer "fizz", inclusive quando também puder dizer "busao". Nos demais casos, é só mostrar o contador de veículos.]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0, 0, 10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>
veiculos, onibus = 0, 0
for c in input():
    veiculos += 1
    if c == 'O':
        onibus += 1

    if 0 == (veiculos % 3):
        if c == 'O' and 0 == (onibus % 5):
            print('fizzbusao')
        else:
            print('fizz')
    elif c == 'O' and 0 == (onibus % 5):
        print('busao')
    else:
        print(veiculos)
</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs>1</cputimelimitsecs>
    <memlimitmb>256</memlimitmb>
    <sandboxparams><![CDATA[{"linkargs":["-lm"]}]]></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>CCCCC
</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
5
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>COCO
</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>OOOOOCCCCCOOOOO
</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
busao
fizz
7
8
fizz
10
11
fizz
13
14
fizzbusao
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>OOCCCCCOCOOOCCCOOOOCCCCCCCOOOOOCCCCOCOCCCCCOOCOCOOOCOOOCCOOCCOOOCCOCCCOCOOCCCCOOCCOCOCCCOCOCOCCCOCOOCOCCCCOCOOOOOCCOCCCCOCCCCOOOCOCCOOOOCCCOCOCOOCOCCOOCCOOOCOCOCOCCOCOCCCOCOOOCOCOCCCOCCOOCCOOCCOCOOCOCCOCCOCOOOCOCOCOCCCOCCOCCCCOCCCCCCCCC
</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
5
fizz
7
8
fizz
10
busao
fizz
13
14
fizz
16
17
fizz
busao
20
fizz
22
23
fizz
25
26
fizz
28
29
fizz
busao
32
fizz
34
35
fizz
37
38
fizz
40
41
fizz
43
44
fizz
46
busao
fizz
49
50
fizz
52
53
fizzbusao
55
56
fizz
58
59
fizz
61
62
fizzbusao
64
65
fizz
67
68
fizz
70
71
fizz
73
busao
fizz
76
77
fizz
79
80
fizz
82
83
fizz
85
86
fizz
88
busao
fizz
91
92
fizz
94
95
fizz
97
98
fizz
busao
101
fizz
103
104
fizz
106
107
fizz
109
110
fizzbusao
112
113
fizz
115
116
fizz
118
119
fizz
121
122
fizz
124
125
fizzbusao
127
128
fizz
130
131
fizz
133
busao
fizz
136
137
fizz
139
140
fizz
142
143
fizzbusao
145
146
fizz
148
149
fizz
151
152
fizz
busao
155
fizz
157
158
fizz
160
161
fizzbusao
163
164
fizz
166
167
fizz
169
170
fizz
172
173
fizzbusao
175
176
fizz
178
179
fizz
181
182
fizz
184
185
fizzbusao
187
188
fizz
190
191
fizz
193
194
fizz
busao
197
fizz
199
200
fizz
202
203
fizz
205
206
fizzbusao
208
209
fizz
211
212
fizz
214
busao
fizz
217
218
fizz
220
221
fizz
223
224
fizz
226
227
fizz
229
230
fizz
232
233
fizz
235
236
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>CCCCCCCCOCOOOCCCCOOOCOCOOCOOCOOOOCCOCOCCOCCCOOCCCCCCOCCOOOCOCCCCCOCCOCOOCOCCCOOOCCCOOOCOOOCOOCOCCOCCCCOCOOCOOCCCOCOCCCCOOCCCCOCCOCCCOCCOCCCOCOOCOOOOCOCOCOCOOCCOCOCOOOCOOOO
</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
5
fizz
7
8
fizz
10
11
fizz
13
14
fizz
16
17
fizzbusao
19
20
fizz
22
23
fizz
busao
26
fizz
28
29
fizz
31
busao
fizz
34
35
fizz
37
38
fizz
40
41
fizz
43
44
fizzbusao
46
47
fizz
49
50
fizz
52
53
fizz
55
56
fizz
busao
59
fizz
61
62
fizz
64
65
fizz
67
68
fizz
70
71
fizzbusao
73
74
fizz
76
77
fizz
79
80
fizz
82
83
fizzbusao
85
86
fizz
88
89
fizzbusao
91
92
fizz
94
95
fizz
97
98
fizz
100
101
fizz
busao
104
fizz
106
107
fizz
109
110
fizz
112
busao
fizz
115
116
fizz
118
119
fizz
121
122
fizz
124
125
fizz
127
128
fizzbusao
130
131
fizz
133
134
fizz
136
137
fizz
139
140
fizz
142
busao
fizz
145
146
fizz
148
149
fizzbusao
151
152
fizz
154
155
fizz
157
158
fizz
busao
161
fizz
163
164
fizz
166
167
fizzbusao
169
170
fizz
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>OCCCOCCOOOOCCOCOOOCOOCOCCOOCCCOCCOCCOCCCOOCOCCCCOOOCCCCCCCCOOOCCCCCCOCCCOOOOCOCOCOCOCOOCCCCCOOOOOCOOCOOOOOCOCOCCOCCCOCOCCCOOOCCOOOCCCCOOOOCCOOOOCCOCOOOCOOCOOCCCCCOCOOOCCCCOCOOOCCOOOOOOCOOOOOOOCOCOCOOOOCCOCCOOOOOOCOOCCCOCOCOCCCOOCOCOOCOCOOCOCCCOOOCCCCOOOOOCOOCCOOCCOOOOOOCOCCOOCCCOCOCOOOCOOOCOCOCOCOOCCOOCOOOOCOCOOOCOOOOCCCCOCOCCOOOOCCCCCCCCOOCOOOOOCOOCCOCCCCOCOCCOCOOCOCOCCCCOCOCCCOCCCCCCOCCCCOOCOOOCOCOOCOOOCOCOCCCOOCCOCOCCOOOOOCOOCOCCCCCCOOOOCCCCOCOCCOOCOCCOOOOCOCCCCOCCCOOOCOOCCCOOOOOOCOCOCOOOOCCCCCCCOOCCOOCCOOCOCCCOCCCCOCCCOOCOOCOOOCCOOCCOOOCCCOOCCCOOCOOCCCCOCCOCCOCOCCCOCCCOOCCCCOOOCOOCCOOOOCOOCOCOOCOCOOCCCOCCOOOCOOOOCCCOCCCOCOCCOOCCOOCCCCOCOCOOCCCCOOCOOOOCCCOCCCOCCCCCCOCCOOCOCOOCOCCOOCCCOCOCOOCOCCOOCOOCCCCCCCOOOCCOCCOOCCOCCCOCCCOCCOOCCOOCOCOOCOOCOOOOOCCCCCCOOCCOOOOCCCCOOOCCOCCCOCCCCOCOOCCOOCCCOCOOOOOOCOOOCCCCCOCC
</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
5
fizz
7
8
fizz
busao
11
fizz
13
14
fizz
16
17
fizzbusao
19
20
fizz
22
23
fizz
25
26
fizzbusao
28
29
fizz
31
32
fizz
34
35
fizz
37
38
fizz
40
41
fizzbusao
43
44
fizz
46
47
fizz
49
50
fizz
52
53
fizz
55
56
fizz
58
59
fizzbusao
61
62
fizz
64
65
fizz
67
68
fizz
70
71
fizz
73
busao
fizz
76
77
fizz
79
80
fizz
busao
83
fizz
85
86
fizz
88
89
fizz
91
92
fizz
busao
95
fizz
97
98
fizz
busao
101
fizz
103
104
fizz
busao
107
fizz
109
110
fizz
112
113
fizz
115
116
fizz
118
busao
fizz
121
122
fizz
124
125
fizz
127
128
fizzbusao
130
131
fizz
133
134
fizz
136
137
fizzbusao
139
140
fizz
142
143
fizz
145
146
fizzbusao
148
149
fizz
151
152
fizz
busao
155
fizz
157
158
fizz
160
161
fizz
163
164
fizz
busao
167
fizz
169
170
fizz
172
173
fizz
175
busao
fizz
178
179
fizz
181
182
fizzbusao
184
185
fizz
187
188
fizzbusao
190
191
fizz
193
194
fizz
busao
197
fizz
199
200
fizz
202
203
fizzbusao
205
206
fizz
208
209
fizz
busao
212
fizz
214
215
fizz
217
218
fizz
220
busao
fizz
223
224
fizz
226
227
fizz
229
230
fizz
busao
233
fizz
235
236
fizz
238
239
fizzbusao
241
242
fizz
244
245
fizz
247
248
fizz
250
251
fizzbusao
253
254
fizz
256
257
fizzbusao
259
260
fizz
262
263
fizz
265
266
fizzbusao
268
269
fizz
271
272
fizz
274
busao
fizz
277
278
fizz
280
281
fizz
283
284
fizzbusao
286
287
fizz
289
290
fizz
busao
293
fizz
295
296
fizz
298
299
fizz
301
busao
fizz
304
305
fizz
307
busao
fizz
310
311
fizz
313
314
fizz
busao
317
fizz
319
320
fizz
322
323
fizz
325
busao
fizz
328
329
fizz
331
332
fizz
334
335
fizz
337
338
fizz
340
busao
fizz
343
344
fizz
346
busao
fizz
349
350
fizz
352
353
fizz
355
356
fizz
358
busao
fizz
361
362
fizz
364
365
fizz
367
368
fizzbusao
370
371
fizz
373
374
fizz
376
377
fizz
379
380
fizz
382
383
fizz
385
386
fizz
388
busao
fizz
391
392
fizz
394
395
fizz
397
398
fizzbusao
400
401
fizz
403
404
fizz
406
busao
fizz
409
410
fizz
412
413
fizz
415
416
fizzbusao
418
419
fizz
421
422
fizz
424
425
fizz
busao
428
fizz
430
431
fizz
433
busao
fizz
436
437
fizz
439
440
fizz
442
443
fizz
445
446
fizz
448
busao
fizz
451
452
fizz
454
455
fizz
457
458
fizz
busao
461
fizz
463
464
fizz
466
467
fizz
469
busao
fizz
472
473
fizz
475
476
fizz
478
busao
fizz
481
482
fizz
484
485
fizz
busao
488
fizz
490
491
fizz
493
494
fizzbusao
496
497
fizz
499
500
fizz
502
503
fizz
505
506
fizz
508
busao
fizz
511
512
fizz
514
515
fizz
517
518
fizz
busao
521
fizz
523
524
fizz
526
527
fizz
529
530
fizz
532
busao
fizz
535
536
fizz
538
539
fizz
busao
542
fizz
544
545
fizz
547
548
fizz
550
busao
fizz
553
554
fizz
556
557
fizz
559
560
fizz
562
563
fizzbusao
565
566
fizz
568
569
fizz
571
572
fizz
574
575
fizz
577
578
fizz
busao
581
fizz
583
584
fizz
586
587
fizz
589
busao
fizz
592
593
fizz
595
596
fizzbusao
598
599
fizz
601
602
fizz
604
busao
fizz
607
608
fizz
610
611
fizz
613
614
fizz
616
busao
fizz
619
620
fizz
622
busao
fizz
625
626
fizz
628
629
fizz
631
632
fizz
634
635
fizz
busao
638
fizz
640
641
fizz
643
644
fizz
646
647
fizz
busao
650
fizz
652
653
fizz
655
656
fizz
658
659
fizzbusao
661
662
fizz
664
665
fizz
667
668
fizz
670
busao
fizz
673
674
fizz
676
677
fizz
679
680
fizz
682
683
fizz
685
busao
fizz
688
689
fizz
691
692
fizz
694
695
fizz
busao
698
fizz
700
701
fizz
703
704
fizz
706
busao
fizz
709
710
fizz
712
713
fizz
715
716
fizz
718
719
fizzbusao
721
722
fizz
724
725
fizz
727
728
fizz
730
busao
fizz
733
734
fizz
736
737
fizz
739
740
fizz
742
743
fizz
745
busao
fizz
748
749
fizz
751
752
fizz
busao
755
fizz
757
758
fizz
busao
761
fizz
763
764
fizz
766
767
fizz
769
770
fizz
772
busao
fizz
775
776
fizz
778
779
fizz
781
busao
fizz
784
785
fizz
787
788
fizz
790
791
fizz
793
794
fizz
796
busao
fizz
799
800
fizz
802
803
fizz
805
806
fizz
busao
809
fizz
811
812
fizz
busao
815
fizz
817
818
fizz
820
821
fizz
823
824
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>CCCCOCCCOOOOOOCCCOCCOOCOCOCOCCOOOCCCOCOOCOOCOOCCOOOOCCOCOCOOCCCCCCOCCOOOCCOOCOOOOOOCCCOOCOOOCOOCCCOCCOCOOOOOCCCCCOCCOOCOCCOOCOOOCOCCCOOOCCCCCCCOCCOOOOCCOOCCOOOOCCOCCCOOCCOCCCOCCOOCCCCCCCOOCCOOOCOOOOCOCOOCCOCOCOOCCOOOCCCCCCOOCOOCOCOOCCCOCCCCOCOCCOCOCCCCCCOOCCCOOOCCOOOOCCCCCCCCOOCCCCCCOOCOCCCCOCCCCOCCCCOOCCCOOCCOCCOOCCOOOOOOOCOCCCOCOOCCCCCOOCOCOOOOOCCCCCOOOCCCCOCOCCOCOOOCOCOOCCOCOCOOCCCOOCOCCCCCCOCOOOOOOCOCOCOOCOOOCOOCCOCOOCOCCCCOCOCCCCOOCCOOOCOOCOCCCCOOOCOOOOOOOOOOOOCOOCOOCOOOOCOCOOCOCOOCCCOOCOCCCOCCCCCOOOCOCOCCOCCOCOOOOOCOOCOOOCOCCOOCCOOOCCCOOCOOOOOCCOOCOOOOOOOOCOOOOOOCCCOCOCCOOOOCCOOCCOOOOOCCOOCOCOCOOOCCCCOOCCCOOOCOOOOCOOOCCOCOCOOOCOCCCOCCOOOCCOCCOCOCCOCOCOOOCOCCCCCOOOCOCOOOCCCOOOOCCOOCOOOOOOOOCCCCOCCCCCCCCOCOCOCCOOOOCOOOOOOOOCCCCOOOCOOCOCCCCOOCOCCCOCCCCOOCCCCOOOCCCOCOCOCCCOOOCCCCOCOCOOOOOOCCOOOOCCOCCCCCCCCCOCCCOOOOCCCCOOOOOOOOCOCCOCCCOOOOOCOCCCCOOCCOCOCCOOOOCOCCCOOOOCCCOCCCOCCOCOCCCOOOOOCOCCCCCCOCOCOOOOOOOOCCCCCCOCOCOCOOCCCOOOOCOOCOOOOOOOCCOOOCOCOCCOOCOOCOCCCOCOCOCOOCCOCCOOCOOOOOOOOC
</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
5
fizz
7
8
fizz
10
11
fizzbusao
13
14
fizz
16
17
fizz
19
20
fizz
busao
23
fizz
25
26
fizz
28
29
fizz
31
busao
fizz
34
35
fizz
37
38
fizz
40
41
fizzbusao
43
44
fizz
46
47
fizz
49
busao
fizz
52
53
fizz
55
56
fizz
58
busao
fizz
61
62
fizz
64
65
fizz
67
68
fizz
70
71
fizzbusao
73
74
fizz
76
77
fizz
79
busao
fizz
82
83
fizz
85
86
fizz
busao
89
fizz
91
92
fizz
94
busao
fizz
97
98
fizz
100
101
fizz
103
104
fizz
busao
107
fizz
109
110
fizz
112
113
fizz
115
116
fizz
busao
119
fizz
121
122
fizz
124
125
fizz
busao
128
fizz
130
131
fizz
133
134
fizz
busao
137
fizz
139
140
fizz
142
143
fizz
145
146
fizz
148
149
fizzbusao
151
152
fizz
154
155
fizz
157
158
fizzbusao
160
161
fizz
163
164
fizz
166
167
fizz
169
170
fizzbusao
172
173
fizz
175
176
fizz
178
179
fizz
181
182
fizz
184
185
fizz
187
busao
fizz
190
191
fizz
193
194
fizz
busao
197
fizz
199
200
fizz
202
busao
fizz
205
206
fizz
208
209
fizz
211
212
fizz
busao
215
fizz
217
218
fizz
220
221
fizz
223
224
fizz
busao
227
fizz
229
230
fizz
232
233
fizz
235
busao
fizz
238
239
fizz
241
242
fizz
244
245
fizz
247
248
fizz
250
251
fizz
253
254
fizzbusao
256
257
fizz
259
260
fizz
262
263
fizz
busao
266
fizz
268
269
fizz
271
272
fizz
274
275
fizz
277
busao
fizz
280
281
fizz
283
284
fizz
286
287
fizz
289
290
fizz
292
293
fizz
295
296
fizz
busao
299
fizz
301
302
fizz
304
305
fizz
307
308
fizz
310
311
fizzbusao
313
314
fizz
316
317
fizz
319
320
fizzbusao
322
323
fizz
325
326
fizzbusao
328
329
fizz
331
332
fizz
334
335
fizz
337
338
fizz
340
busao
fizz
343
344
fizz
346
347
fizzbusao
349
350
fizz
352
353
fizz
355
356
fizz
358
359
fizz
361
busao
fizz
364
365
fizz
367
368
fizz
370
busao
fizz
373
374
fizz
376
377
fizz
379
380
fizzbusao
382
383
fizz
385
386
fizz
388
389
fizz
busao
392
fizz
394
395
fizz
397
398
fizz
400
401
fizz
busao
404
fizz
406
407
fizz
409
410
fizzbusao
412
413
fizz
415
416
fizz
busao
419
fizz
421
422
fizz
424
425
fizz
busao
428
fizz
430
431
fizz
433
434
fizz
436
437
fizz
439
440
fizz
442
busao
fizz
445
446
fizz
448
449
fizzbusao
451
452
fizz
454
455
fizz
457
458
fizz
busao
461
fizz
463
464
fizzbusao
466
467
fizz
469
busao
fizz
472
473
fizz
475
476
fizz
busao
479
fizz
481
482
fizz
484
busao
fizz
487
488
fizz
490
491
fizz
493
494
fizzbusao
496
497
fizz
499
500
fizz
502
503
fizz
505
506
fizz
508
busao
fizz
511
512
fizz
514
515
fizz
517
518
fizz
busao
521
fizz
523
524
fizz
busao
527
fizz
529
530
fizz
532
busao
fizz
535
536
fizz
538
539
fizz
541
542
fizzbusao
544
545
fizz
547
548
fizz
550
551
fizzbusao
553
554
fizz
556
557
fizz
busao
560
fizz
562
563
fizz
busao
566
fizz
568
569
fizz
busao
572
fizz
574
575
fizz
577
578
fizzbusao
580
581
fizz
583
584
fizz
586
busao
fizz
589
590
fizz
592
593
fizz
595
busao
fizz
598
599
fizz
601
602
fizz
busao
605
fizz
607
608
fizz
610
611
fizz
613
614
fizzbusao
616
617
fizz
619
620
fizz
622
623
fizzbusao
625
626
fizz
628
629
fizzbusao
631
632
fizz
634
635
fizz
637
638
fizzbusao
640
641
fizz
643
644
fizz
646
647
fizz
649
busao
fizz
652
653
fizz
655
656
fizz
658
659
fizz
661
busao
fizz
664
665
fizz
667
668
fizz
busao
671
fizz
673
674
fizz
676
677
fizz
679
680
fizz
busao
683
fizz
685
686
fizz
688
689
fizzbusao
691
692
fizz
694
695
fizz
697
busao
fizz
700
701
fizz
busao
704
fizz
706
707
fizz
709
710
fizz
712
713
fizz
715
716
fizz
718
719
fizz
721
busao
fizz
724
725
fizz
727
728
fizz
busao
731
fizz
733
734
fizzbusao
736
737
fizz
739
740
fizz
742
743
fizzbusao
745
746
fizz
748
749
fizz
751
752
fizz
754
busao
fizz
757
758
fizz
760
761
fizz
763
764
fizz
766
767
fizz
769
770
fizz
busao
773
fizz
775
776
fizz
778
779
fizz
781
busao
fizz
784
785
fizz
787
788
fizz
790
791
fizz
793
794
fizzbusao
796
797
fizz
799
800
fizzbusao
802
803
fizz
805
806
fizz
busao
809
fizz
811
812
fizz
814
815
fizz
817
818
fizz
820
821
fizz
823
824
fizz
826
busao
fizz
829
830
fizz
832
833
fizz
835
busao
fizz
838
839
fizz
841
busao
fizz
844
845
fizz
847
848
fizz
850
851
fizzbusao
853
854
fizz
856
857
fizz
859
860
fizz
862
863
fizzbusao
865
866
fizz
868
869
fizz
871
busao
fizz
874
875
fizz
877
878
fizz
880
busao
fizz
883
884
fizz
886
887
fizz
889
890
fizz
892
893
fizz
895
896
fizz
busao
899
fizz
901
902
fizz
busao
905
fizz
907
908
fizz
910
911
fizz
913
914
fizz
916
busao
fizz
919
920
fizz
busao
923
fizz
925
926
fizz
928
929
fizz
931
932
fizz
934
935
fizzbusao
937
938
fizz
940
941
fizz
943
944
fizzbusao
946
947
fizz
949
950
fizzbusao
952
953
fizz
955
956
fizz
busao
959
fizz
961
962
fizz
964
965
fizz
busao
968
fizz
970
971
fizz
973
974
fizz
976
977
fizzbusao
979
980
fizz
982
983
fizz
985
986
fizz
988
busao
fizz
991
992
fizz
994
busao
fizz
997
998
fizz
1000
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>OCOCCOCCOOOCCOOOCOCOOCOOOOOCOOCCCOOOOCCCOCOCOOCCOOOOOOCOCOCCCOCCOCCCCCOOOOCOOOCOCOOOCCOCCOOCCOCCCOCCCOCOOOCOCOOOCOOOOCOCCOCOOCCCOCOCOOOOOOOOCCCOOOOOOCOCOCCCCOCCCOCCOOOOOOCOCOOOOOOOOCCOOCOCOOOCCCOCCCOOOCOCCCOOOOOOCOOCOOOOOOOOOOOCCCCCCOOCCCCOOCOCCOCOOCCCCCOCOCOOCOCOOCCCCCOCOOCCOOOOCCOCCCCOOCCCCCCCCOCOOCCOCCCOCCCCOOOCCCCCCCOOCOCCCCOCCCCCCOCCOCCOOCOCOCCOCOCCCCCOOCOCOOCOOOCOCOCCOCOOOOOOOCOOOCOCCOCCCCCOCCOCOOCOCCOOCCOCOOOOCOOOCOCOCCOOCCOCCCOCCOCCCCOOCOOOOCCOCOCOCCCOOOOCOCCCOOCCOOCOOCCCOCOCOCOOOCOCOOCOOOCOOOCOCOOCCCOCCCCOOCCOCCOOOCOOCCCCCOCCCCOCCCCOOOCOOOCCCOOCOOCOCCOOCOOOCCCOCOOOCOOOCCOCOOCCOCOCOCCCCCOOOCOCCOOCOOCCCCOOCCCCOCOCOOOOCCOOOCOOCCOCOCOOOCOOCCCCOCCCOCOOCOOCCCCOCCOOOCOCOOCCCCCOOOOOOCOOCCOCOCCCCOCOOCCOCOOOCOOOCCOOOCCCCOOOOOCCCOCOOOOOOCCOCCOCOOCCCCCOOCOOCOOCOOCOOCOOOOCCCCCCOCOOCCCOOOCCOOCCCCOOCOCOOCOCCCCOOOCCOCCCCOCOOOCOOOCCCCCCOCOOOCOCOCCCOOCOCCOOOCOCOOOOCOOOCCOOOOOOOOOOOOOOCOOC
</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
5
fizz
7
8
fizz
busao
11
fizz
13
14
fizz
16
17
fizzbusao
19
20
fizz
22
23
fizz
busao
26
fizz
28
29
fizz
31
32
fizz
busao
35
fizz
37
38
fizz
40
41
fizz
busao
44
fizz
46
47
fizz
49
50
fizzbusao
52
53
fizz
55
56
fizz
busao
59
fizz
61
62
fizz
64
65
fizz
67
68
fizz
70
71
fizz
busao
74
fizz
76
77
fizz
79
busao
fizz
82
83
fizz
85
86
fizz
88
89
fizzbusao
91
92
fizz
94
95
fizz
97
98
fizz
100
101
fizz
103
busao
fizz
106
107
fizz
109
110
fizzbusao
112
113
fizz
115
116
fizzbusao
118
119
fizz
121
122
fizz
124
125
fizz
127
128
fizzbusao
130
131
fizz
133
134
fizz
busao
137
fizz
139
140
fizz
142
143
fizzbusao
145
146
fizz
148
busao
fizz
151
152
fizz
154
155
fizz
157
158
fizz
160
161
fizz
163
164
fizzbusao
166
167
fizz
169
busao
fizz
172
173
fizz
175
176
fizzbusao
178
179
fizz
181
182
fizz
busao
185
fizz
187
188
fizz
190
busao
fizz
193
194
fizz
196
197
fizz
199
200
fizz
202
busao
fizz
205
206
fizz
208
209
fizz
busao
212
fizz
214
215
fizz
217
busao
fizz
220
221
fizz
busao
224
fizz
226
227
fizz
229
230
fizz
232
233
fizzbusao
235
236
fizz
238
239
fizz
241
242
fizz
244
245
fizzbusao
247
248
fizz
250
251
fizz
253
254
fizz
256
257
fizz
busao
260
fizz
262
263
fizz
265
266
fizz
268
269
fizz
busao
272
fizz
274
275
fizz
277
278
fizzbusao
280
281
fizz
283
284
fizz
286
287
fizz
289
290
fizz
292
293
fizz
295
296
fizz
busao
299
fizz
301
302
fizz
304
305
fizz
307
308
fizz
310
311
fizz
busao
314
fizz
316
317
fizz
319
320
fizz
322
323
fizz
325
busao
fizz
328
329
fizz
331
332
fizz
334
335
fizz
337
338
fizz
340
341
fizz
343
344
fizzbusao
346
347
fizz
349
350
fizz
352
353
fizz
355
356
fizz
358
359
fizzbusao
361
362
fizz
364
365
fizz
367
busao
fizz
370
371
fizz
373
374
fizz
376
busao
fizz
379
380
fizz
382
busao
fizz
385
386
fizz
388
busao
fizz
391
392
fizz
394
395
fizz
397
398
fizz
400
401
fizz
403
404
fizzbusao
406
407
fizz
409
410
fizz
412
413
fizz
busao
416
fizz
418
419
fizz
421
busao
fizz
424
425
fizz
427
428
fizz
430
busao
fizz
433
434
fizz
436
437
fizz
439
440
fizz
442
443
fizz
445
446
fizzbusao
448
449
fizz
451
452
fizzbusao
454
455
fizz
457
458
fizz
460
461
fizz
463
464
fizzbusao
466
467
fizz
469
470
fizz
472
473
fizzbusao
475
476
fizz
478
479
fizz
481
482
fizz
484
busao
fizz
487
488
fizz
490
491
fizz
busao
494
fizz
496
497
fizz
499
500
fizzbusao
502
503
fizz
505
506
fizz
busao
509
fizz
511
512
fizz
514
515
fizz
517
518
fizz
520
busao
fizz
523
524
fizz
526
527
fizz
529
530
fizzbusao
532
533
fizz
535
536
fizz
538
539
fizz
541
542
fizz
544
545
fizz
547
548
fizzbusao
550
551
fizz
553
554
fizz
556
557
fizzbusao
559
560
fizz
562
563
fizz
565
566
fizzbusao
568
569
fizz
571
572
fizz
574
575
fizzbusao
577
578
fizz
580
581
fizz
busao
584
fizz
586
587
fizz
589
590
fizz
592
busao
fizz
595
596
fizz
598
599
fizz
601
602
fizz
604
busao
fizz
607
608
fizz
610
611
fizz
613
busao
fizz
616
617
fizz
619
620
fizz
622
623
fizz
625
626
fizz
628
busao
fizz
631
632
fizz
634
635
fizzbusao
637
638
fizz
640
641
fizz
643
644
fizzbusao
646
647
fizz
649
650
fizz
busao
653
fizz
655
656
fizz
658
659
fizz
661
662
fizz
664
665
fizzbusao
667
668
fizz
670
671
fizz
673
674
fizz
676
busao
fizz
679
680
fizz
682
683
fizz
685
686
fizz
688
busao
fizz
691
692
fizz
694
busao
fizz
697
698
fizz
700
701
fizz
703
704
fizz
706
707
fizzbusao
709
710
fizz
712
713
fizz
715
busao
fizz
718
719
fizz
721
722
fizz
busao
725
fizz
727
728
fizz
730
731
fizz
busao
734
fizz
736
737
fizz
739
740
fizz
busao
743
fizz
745
746
fizz
748
749
fizz
busao
752
fizz
754
755
fizz
757
758
fizz
760
761
fizz
busao
764
fizz
766
767
fizz
769
busao
fizz
772
773
fizz
775
776
fizzbusao
778
779
fizz
781
782
fizz
784
785
fizz
787
788
fizz
790
791
fizzbusao
793
794
fizz
796
797
fizz
799
800
fizz
802
busao
fizz
805
806
fizz
808
809
fizz
busao
812
fizz
814
815
fizz
817
818
fizz
820
821
fizz
823
824
fizz
busao
827
fizz
829
830
fizz
832
busao
fizz
835
836
fizz
838
839
fizz
841
842
fizz
844
busao
fizz
847
848
fizz
850
851
fizz
853
854
fizz
busao
857
fizz
859
860
fizz
862
863
fizz
busao
866
fizz
868
869
fizz
busao
872
fizz
874
875
fizz
877
busao
fizz
880
881
fizz
busao
884
fizz
886
887
fizzbusao
889
890
fizz
892
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>COOCOCCOOCOCCOCCOOOCOCOOOCOOOOOCCOCOCOOOOCOOOCOOOOCOOCOCCOOOOCCOCCCOCCCCOCOOCOOCOOCCCOCOCOCCCOOCCCCOCOCCCOOCOOCOOOCOCCCOOCCCOOOCOOCCOOCCOOCOCOCOCCOOCCOOOCCOOCOOOCOCCOCCOOCOCCOOCCOCCOCCCCCOOOCCOCCOCCOOCOCOOCOOOOOCCOCCCOCCOCOCCOCOCCCOOOOOCCOCCCCOOOCOOOCCCCCOCCCCCCCCOOCCCOCCCOOOCCOOOCOOOCOOCOOCCCOCCCOCOOOCOOCOCCCOOCCCCOOCCOCCOCOOOOCOOOCCOOOCOOOCCOOCOCCOCCOOCCOCCOCCOOCCCCOCCCCCCCOCOOCCOCOCCCOCCCOOCOOOCOOCCCOCCOOCOOCCCOCCOOOOOOCOOOCOOOOCCCOCOCCOCOOCCCCOOOOOOCOCCOOCOOOOCCOOCOCCCCCCCOOCOOOOOCOCCCCOOCCCCOOCOCCCCOOOCCCCOOOOOOOOOOCCOOCCCOOOCCOCCOCCOOOCOCOOCCCOOCCCOOOOOOOOOOOOCOCCCOOCOOCOOCCOCCCOOOCOCCOCCCCOCOCOOCCCCOOOCCCCCOCCCOOCOCCCCCOOOOOCCOOCOCOCCCCOOOOOCOCCCCCOOOOOOOOOOCCOCOOOCCCOOCOCCOCCCCOOOCOOCOOOCOCOOOCOCCOOCOOOCCOOOOCCCCOCOOOOCOOCCCCOC
</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
5
fizz
7
8
fizzbusao
10
11
fizz
13
14
fizz
16
17
fizz
busao
20
fizz
22
23
fizz
25
26
fizzbusao
28
29
fizz
31
32
fizz
busao
35
fizz
37
38
fizz
40
busao
fizz
43
44
fizz
46
47
fizzbusao
49
50
fizz
52
53
fizz
busao
56
fizz
58
59
fizz
61
62
fizz
busao
65
fizz
67
68
fizz
70
71
fizz
73
74
fizz
76
77
fizzbusao
79
80
fizz
82
83
fizz
85
86
fizz
busao
89
fizz
91
92
fizz
94
95
fizz
97
98
fizz
100
101
fizzbusao
103
104
fizz
106
107
fizz
109
110
fizz
busao
113
fizz
115
116
fizz
118
119
fizz
busao
122
fizz
124
125
fizz
127
128
fizz
busao
131
fizz
133
134
fizz
136
137
fizz
139
busao
fizz
142
143
fizz
145
146
fizz
148
149
fizz
busao
152
fizz
154
155
fizz
157
158
fizzbusao
160
161
fizz
163
164
fizz
166
167
fizz
busao
170
fizz
172
173
fizz
175
176
fizz
178
busao
fizz
181
182
fizz
184
185
fizz
187
188
fizz
190
191
fizz
busao
194
fizz
196
197
fizz
199
200
fizz
202
203
fizzbusao
205
206
fizz
208
209
fizzbusao
211
212
fizz
214
215
fizz
217
218
fizz
220
221
fizz
busao
224
fizz
226
227
fizz
229
230
fizz
232
233
fizzbusao
235
236
fizz
238
239
fizz
241
242
fizz
244
busao
fizz
247
248
fizz
250
251
fizz
253
254
fizz
busao
257
fizz
259
260
fizz
262
263
fizz
265
266
fizz
268
269
fizz
271
272
fizz
274
busao
fizz
277
278
fizz
280
281
fizz
busao
284
fizz
286
287
fizz
289
busao
fizz
292
293
fizz
295
296
fizz
298
299
fizz
301
busao
fizz
304
305
fizz
307
308
fizz
310
311
fizzbusao
313
314
fizz
316
317
fizz
319
320
fizz
322
323
fizz
busao
326
fizz
328
329
fizz
331
busao
fizz
334
335
fizz
337
338
fizzbusao
340
341
fizz
343
344
fizz
346
busao
fizz
349
350
fizz
352
353
fizz
355
356
fizz
358
busao
fizz
361
362
fizz
364
365
fizz
367
368
fizz
370
371
fizz
373
374
fizz
376
377
fizz
busao
380
fizz
382
383
fizz
385
386
fizz
388
389
fizz
busao
392
fizz
394
395
fizz
397
398
fizz
busao
401
fizz
403
404
fizz
406
407
fizz
409
410
fizzbusao
412
413
fizz
415
416
fizz
418
419
fizz
421
busao
fizz
424
425
fizz
427
busao
fizz
430
431
fizz
433
busao
fizz
436
437
fizz
439
440
fizz
442
443
fizz
445
busao
fizz
448
449
fizz
451
452
fizz
454
busao
fizz
457
458
fizz
460
461
fizz
busao
464
fizz
466
467
fizz
469
470
fizzbusao
472
473
fizz
475
476
fizz
478
479
fizz
481
482
fizz
484
busao
fizz
487
488
fizz
490
busao
fizz
493
494
fizz
496
497
fizz
499
500
fizz
502
503
fizz
busao
506
fizz
508
509
fizz
511
512
fizz
514
515
fizz
517
busao
fizz
520
521
fizz
busao
524
fizz
526
527
fizz
529
busao
fizz
532
533
fizz
535
536
fizz
538
539
fizz
541
busao
fizz
544
545
fizz
547
548
fizz
550
busao
fizz
553
554
fizz
556
557
fizz
559
560
fizz
busao
563
fizz
565
566
fizzbusao
568
569
fizz
571
busao
fizz
574
575
fizz
577
578
fizz
580
581
fizzbusao
583
584
fizz
586
587
fizz
589
590
fizz
592
busao
fizz
595
596
fizz
598
599
fizz
601
602
fizz
604
605
fizzbusao
607
608
fizz
610
611
fizz
613
614
fizz
busao
617
fizz
619
620
fizz
622
623
fizz
625
626
fizz
628
629
fizz
631
632
fizz
634
busao
fizz
637
638
fizz
640
641
fizzbusao
643
644
fizz
646
647
fizz
649
650
fizz
652
busao
fizz
655
656
fizz
658
659
fizz
661
662
fizz
busao
665
fizz
667
668
fizzbusao
670
671
fizz
673
674
fizz
busao
677
fizz
679
680
fizz
682
683
fizz
busao
686
fizz
688
689
fizz
691
692
fizz
694
695
fizz
busao
698
fizz
700
701
fizz
703
busao
fizz
706
707
fizz
709
710
fizz
busao
713
fizz
715
716
fizz
718
719
fizzbusao
721
722
fizz
724
725
fizz
727
728
fizz
730
busao
fizz
733
734
fizz
736
737
fizzbusao
739
740
fizz
742
743
fizz
745
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>OOCCOOCOOOCOCOCOOOOOOOCCCCCOOCCOOOCCCOOCCCCCOOCOCCCCCCCOCCCCOCCCCOCOCCOCCOCOCCOCCCCOCCCCOCOOOOCCOCCCCCCOCCOCCCCOCOCCCOOCCOCOCCOOOOOOCCOCOCOCCOOOOCOCCOCCOCCCOOCCCOOCOOCOCCCOOOCCCOOCCOCOCCCCOOCOCOOCCCCOOOOOOCCOCOCCOOOCOOOOOCCCCCCOCOOCCCOOCOCCCOOCCCCOCCCCCOCCOCOOOOOOCCCOCCOCCOCOCCCOCOOOCOOCCOOOOCCCOCCOOCCCOOOCCOCCOOOCCOOOOCCOOCOOOOOOOCCOOOOCOOOOOCOOCOOCOOOCCOOOCCOCOCCOCOOOCCCOOOCCOOCOCOOCOOCOCCOOOCOCOOOCCOOCOCOOCOCOOCOCOCCCOCCOCCCOOOCCOOCCOOCCOCCOOOOCCOOOOOCCCCOCOOOCOOOCOCOCCOCCOCOCCCOCOOOOCOOOCCCCCCCCCOOOCCOOOCOOCCCOOCCCOCOCCOCCCOCCOCOCCCCCOCCCOOCOCCCOOOOOOOOOCCCOOCOCOOCOOOOCCOOCCCOOCCOCCCCCCOCCOCOCOCCCOCCOOCOOCOOOOCOOCCCCOCOOCCCOOOOOOCOCOOCOCOCCOCCCCCOCOOOCCCOOCOOOOOOCCCCCCCCOCOCOOCOOOOCOOCCCOCCOOOOCCCOOOOCOOO
</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
5
fizz
7
busao
fizz
10
11
fizz
13
14
fizz
busao
17
fizz
19
20
fizzbusao
22
23
fizz
25
26
fizz
28
29
fizz
31
32
fizzbusao
34
35
fizz
37
38
fizz
40
41
fizz
43
44
fizz
busao
47
fizz
49
50
fizz
52
53
fizz
55
56
fizz
58
59
fizz
61
62
fizz
64
65
fizz
67
busao
fizz
70
71
fizz
73
74
fizz
76
77
fizz
79
80
fizz
82
83
fizzbusao
85
86
fizz
88
89
fizz
91
92
fizz
busao
95
fizz
97
98
fizz
100
101
fizz
103
104
fizz
106
107
fizz
109
110
fizz
112
113
fizzbusao
115
116
fizz
118
119
fizz
121
122
fizz
124
125
fizz
busao
128
fizz
130
131
fizzbusao
133
134
fizz
136
137
fizz
139
140
fizz
142
busao
fizz
145
146
fizz
148
149
fizz
151
152
fizzbusao
154
155
fizz
157
158
fizz
160
161
fizz
163
164
fizzbusao
166
167
fizz
169
170
fizz
172
173
fizzbusao
175
176
fizz
178
179
fizz
181
182
fizz
184
185
fizz
187
188
fizzbusao
190
191
fizz
193
194
fizz
196
197
fizz
199
busao
fizz
202
203
fizz
busao
206
fizz
208
209
fizz
211
212
fizz
214
busao
fizz
217
218
fizz
220
busao
fizz
223
224
fizz
226
227
fizz
229
230
fizz
232
233
fizz
235
busao
fizz
238
239
fizz
241
242
fizz
244
245
fizz
247
248
fizz
250
251
fizz
253
busao
fizz
256
257
fizz
259
260
fizz
busao
263
fizz
265
266
fizz
268
269
fizz
271
272
fizz
busao
275
fizz
277
278
fizz
280
281
fizz
283
busao
fizz
286
287
fizz
289
290
fizz
busao
293
fizz
295
296
fizz
298
299
fizz
301
302
fizz
304
busao
fizz
307
308
fizz
310
311
fizz
313
busao
fizz
316
317
fizz
319
320
fizzbusao
322
323
fizz
325
326
fizz
328
busao
fizz
331
332
fizz
334
335
fizzbusao
337
338
fizz
340
341
fizzbusao
343
344
fizz
346
347
fizzbusao
349
350
fizz
352
353
fizz
busao
356
fizz
358
359
fizz
361
362
fizz
364
busao
fizz
367
368
fizz
370
371
fizz
373
374
fizz
busao
377
fizz
379
380
fizz
382
383
fizzbusao
385
386
fizz
388
389
fizz
391
busao
fizz
394
395
fizz
397
398
fizz
400
busao
fizz
403
404
fizz
406
407
fizz
busao
410
fizz
412
413
fizz
415
416
fizzbusao
418
419
fizz
421
422
fizz
424
425
fizz
427
428
fizz
430
431
fizzbusao
433
434
fizz
436
437
fizz
439
440
fizzbusao
442
443
fizz
445
446
fizz
448
449
fizzbusao
451
452
fizz
454
455
fizz
busao
458
fizz
460
461
fizz
463
464
fizz
466
busao
fizz
469
470
fizz
472
473
fizz
busao
476
fizz
478
479
fizz
481
482
fizz
484
485
fizz
487
488
fizzbusao
490
491
fizz
493
494
fizzbusao
496
497
fizz
499
500
fizz
502
503
fizz
505
506
fizz
508
509
fizz
busao
512
fizz
514
515
fizz
517
518
fizz
busao
521
fizz
523
524
fizz
526
527
fizz
529
530
fizz
532
533
fizzbusao
535
536
fizz
538
539
fizz
541
542
fizz
544
545
fizz
547
548
fizz
busao
551
fizz
553
554
fizz
556
557
fizz
busao
560
fizz
562
563
fizzbusao
565
566
fizz
568
569
fizz
571
572
fizz
busao
575
fizz
577
578
fizz
580
581
fizzbusao
583
584
fizz
586
587
fizz
589
590
fizz
592
593
fizz
595
596
fizz
busao
599
fizz
601
602
fizz
604
605
fizz
607
608
fizz
610
611
fizzbusao
613
614
fizz
616
617
fizz
busao
620
fizz
622
623
fizz
625
626
fizz
628
busao
fizz
631
632
fizz
634
635
fizz
637
busao
fizz
640
641
fizz
643
644
fizzbusao
646
647
fizz
649
650
fizz
652
653
fizz
655
656
fizz
658
busao
fizz
661
662
fizz
664
665
fizz
667
busao
fizz
670
671
fizz
673
busao
fizz
676
677
fizz
679
680
fizz
682
683
fizz
685
686
fizz
688
busao
fizz
691
692
fizz
694
695
fizzbusao
697
698
fizz
700
701
fizz
703
704
fizz
busao
707
fizz
709
710
fizz
712
713
fizzbusao
715
716
fizz
718
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>OOCCCOCCOOOCCOOOCOCCCOCCOOCOCCOOCCCCCCCOOCCOCCOCOCOCCCOCCCOCOCCOOOOCCOCCOCCOCCCCCOCOCOOCOOCOCCOCCOOOCCOCOOOOCCOOOOCCOOOOCCCOOCCOOCOCOCCCCOCCOCOOOOOCOOOCCOCOOCCCOOCCCCOOCCOCOCCCOOCOOOCOOCOCCOCOOCOOCOCOOCOOCCCCOOOCOOCOCCCOOCCOOCOCCOOCCCOOOCCOCOOCCCCOOOOCCCCCCCOOCCOOCCOCOOOCCCOOCCCCCOCCCCOOCOCCCOOCOCOOCOOCOCOCCOCCCOCCCOCCCCOCCCCOOCCOOOOCCCCOCCCOOOOCOCOCCOOOOCCOCOCCCOCCOCCCCOOOCCOOOCOCCOOCOCCOCCCCCCCOCOCCCOCCOOCCCCCCOCOOCCCOCCCCOOOCCOCOCOOOOOOOCCOOCCCCOOOOCCOOOCCOOCOOCOOCCCOOCCOCCOOOOO
</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
5
fizz
7
8
fizz
busao
11
fizz
13
14
fizz
16
17
fizzbusao
19
20
fizz
22
23
fizz
25
26
fizz
28
29
fizz
busao
32
fizz
34
35
fizz
37
38
fizz
40
41
fizz
43
44
fizz
46
busao
fizz
49
50
fizz
52
53
fizz
55
56
fizz
58
59
fizz
busao
62
fizz
64
65
fizz
67
68
fizz
busao
71
fizz
73
74
fizz
76
77
fizz
79
80
fizz
82
83
fizz
85
busao
fizz
88
89
fizz
91
92
fizz
94
busao
fizz
97
98
fizz
100
101
fizz
103
104
fizzbusao
106
107
fizz
109
110
fizz
busao
113
fizz
115
116
fizz
118
busao
fizz
121
122
fizz
124
125
fizz
127
128
fizzbusao
130
131
fizz
133
134
fizz
136
137
fizz
139
140
fizz
142
busao
fizz
145
146
fizz
148
busao
fizz
151
152
fizz
154
155
fizz
busao
158
fizz
160
161
fizz
163
164
fizz
166
167
fizz
169
170
fizzbusao
172
173
fizz
175
176
fizz
178
179
fizz
busao
182
fizz
184
185
fizz
187
188
fizz
busao
191
fizz
193
194
fizz
196
197
fizzbusao
199
200
fizz
202
203
fizz
205
206
fizz
208
busao
fizz
211
212
fizz
214
215
fizzbusao
217
218
fizz
220
221
fizz
223
224
fizz
226
busao
fizz
229
230
fizz
232
233
fizz
235
236
fizzbusao
238
239
fizz
241
242
fizz
244
245
fizz
247
248
fizzbusao
250
251
fizz
253
254
fizz
256
257
fizz
259
260
fizz
262
busao
fizz
265
266
fizz
268
269
fizz
busao
272
fizz
274
275
fizz
277
278
fizz
280
281
fizz
283
284
fizz
286
287
fizzbusao
289
290
fizz
292
293
fizz
295
296
fizz
298
busao
fizz
301
302
fizz
304
305
fizz
busao
308
fizz
310
311
fizz
313
314
fizz
316
317
fizz
319
320
fizz
322
323
fizz
325
326
fizz
busao
329
fizz
331
332
fizz
334
busao
fizz
337
338
fizz
340
341
fizz
343
344
fizz
346
busao
fizz
349
350
fizz
352
353
fizz
355
busao
fizz
358
359
fizz
361
362
fizz
364
365
fizz
367
368
fizzbusao
370
371
fizz
373
374
fizz
376
377
fizz
379
busao
fizz
382
383
fizz
385
386
fizz
388
busao
fizz
391
392
fizz
394
395
fizz
397
398
fizz
400
401
fizz
403
404
fizz
406
407
fizz
busao
410
fizz
412
413
fizz
415
416
fizz
418
419
fizz
421
422
fizz
busao
425
fizz
427
428
fizz
430
431
fizz
433
434
fizz
busao
437
fizz
439
440
fizz
busao
443
fizz
445
446
fizz
448
449
fizz
451
452
fizzbusao
454
455
fizz
457
458
fizz
busao
461
fizz
463
464
fizz
466
467
fizzbusao
469
470
fizz
472
473
fizz
475
476
fizz
478
busao
fizz
481
482
fizz
484
485
fizzbusao
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>CCCCCOCCCCOCCOCCOOCOOOCOCOCCOCOCOCOOOCOCCOCOCOOOOOCCCCOCCOOCCOOOOCCCOCOCOOOOCCCOOOCOCOCOCCCOOCOOOOCOCOCCOCCOOCOCOOOOOOCCCOOOOOOCOCOCOOCCCOOCCOCCCCOCOOOOCCOCOOCOCOCCOCCCOOCOCCOCCOOCCCCOOOOOCCOOOOCOOCCOCOOOCCOOCCOCOOCCOOCOCCCOOOCCOOCOOOCCCCOOCCOCOOCCOOCOCCOOOCOCOOOOOOCCCCOCCOCOOOOOCCOOOOOCCOCOOCOCOCOOCCOCCCCOCOOOOCOOCCCCOOOCOOOCCCOOOOCCCCCOCCOOCCOOOCOOCOCCCOOCCOCOCCCCOCCCCCOCCCOOCCCCCOOCCCCOCCCCOCOCCCCCCCOCCOCOOCCOOCOCCOCCCOOOCOCOOOCOCCCCOOCOCOOCCCCOCOOCCCCCOOOOCOCOOOCCCOOOCCCCCCOCOOCOOOO
</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
5
fizz
7
8
fizz
10
11
fizz
13
14
fizz
16
17
fizzbusao
19
20
fizz
22
23
fizz
25
busao
fizz
28
29
fizz
31
32
fizz
34
35
fizzbusao
37
38
fizz
40
41
fizz
43
44
fizz
busao
47
fizz
49
50
fizz
52
53
fizz
busao
56
fizz
58
59
fizz
61
62
fizz
busao
65
fizz
67
68
fizz
70
71
fizz
73
busao
fizz
76
77
fizz
79
80
fizz
busao
83
fizz
85
86
fizz
88
89
fizz
91
92
fizzbusao
94
95
fizz
97
98
fizz
busao
101
fizz
103
104
fizz
106
107
fizz
109
110
fizzbusao
112
113
fizz
115
116
fizzbusao
118
119
fizz
121
122
fizz
124
busao
fizz
127
128
fizz
130
131
fizz
busao
134
fizz
136
137
fizz
139
140
fizz
142
143
fizz
145
146
fizzbusao
148
149
fizz
151
152
fizz
154
busao
fizz
157
158
fizz
160
161
fizz
163
164
fizzbusao
166
167
fizz
169
170
fizz
172
173
fizz
175
176
fizz
busao
179
fizz
181
182
fizz
184
185
fizz
busao
188
fizz
190
191
fizz
193
busao
fizz
196
197
fizz
199
200
fizz
202
busao
fizz
205
206
fizz
208
209
fizz
211
212
fizzbusao
214
215
fizz
217
218
fizz
220
221
fizz
223
busao
fizz
226
227
fizz
229
230
fizz
busao
233
fizz
235
236
fizz
238
239
fizz
241
242
fizzbusao
244
245
fizz
247
248
fizz
250
251
fizzbusao
253
254
fizz
256
257
fizz
259
260
fizzbusao
262
263
fizz
265
busao
fizz
268
269
fizz
271
272
fizz
274
275
fizz
277
busao
fizz
280
281
fizz
283
284
fizzbusao
286
287
fizz
289
290
fizz
292
busao
fizz
295
296
fizz
298
299
fizz
301
302
fizzbusao
304
305
fizz
307
308
fizz
310
311
fizz
busao
314
fizz
316
317
fizz
319
320
fizz
322
busao
fizz
325
326
fizz
328
329
fizz
331
busao
fizz
334
335
fizz
337
338
fizz
340
341
fizz
343
busao
fizz
346
347
fizz
349
350
fizz
busao
353
fizz
355
356
fizz
358
359
fizz
361
362
fizz
busao
365
fizz
367
368
fizz
370
371
fizz
373
374
fizz
376
377
fizz
379
380
fizz
382
383
fizz
385
busao
fizz
388
389
fizz
391
392
fizz
394
395
fizz
397
398
fizz
400
401
fizz
403
404
fizz
406
busao
fizz
409
410
fizz
412
413
fizz
415
416
fizzbusao
418
419
fizz
421
422
fizz
424
425
fizz
427
busao
fizz
430
431
fizz
433
434
fizz
busao
437
fizz
439
440
fizz
442
443
fizz
445
446
fizzbusao
448
449
fizz
451
452
fizz
454
455
fizz
457
458
fizz
460
461
fizzbusao
463
464
fizz
466
467
fizz
busao
470
fizz
472
473
fizz
475
476
fizz
478
479
fizz
481
482
fizzbusao
484
485
fizz
487
488
fizz
busao
491
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>OOCCCCCCOCOCCOCCCOOOOOOCOOOOOOCCOOOOCOOOCOCOOOCCCOCOCCCOOCOCCCOOCOOCOOOCOOCOCOCCOOCOOOCCCCOCCOCOOCOCCOOCOCCOCCCCOOOCOOOOCOOCOCOCOOOOCCCCOCOOCCCOOCCOOCOOOCCCCCCCOCOOCOCOOCCCOOCCOCOCOOCCCOCCOOCCCOOOOOOCCCCCOOOOCCCCOCCCCCOCOOCOOOCOCOOCOOCCOCCOOOCOOCOCCCOCOCOOOOCCOCCOCOOCOOOCOOOOCCCOCOOOOOCCCOOCOCOCOOOCCOOOCOCCOCCOCOCOOOOOCOCCOOCCCCOOOOCOCOOOCCOCOCOCOCOOCOCCOCOOOCCOCCCCCCOOCOOCCOOCCCCCOCOOOCOCCOOCCOCOOOOCCOCCOOCOCOCOCCOOCOCOCCCOOOCCCOCCOCCOOCCCOCOCCCCOOOOOCCCOOCOCCCCOOCOOCOCCCOOCOOOCCOCCCOOCCCOCOOOCCOCOCOCOOCCOOCOCCOOOOOOCOOOOCOOCCCCOOCOOOCCCOOCCCCOOCCOCOOOOOCOOCCCOCOOCCOOCCOOCOOOCOOOOOCCOOCCCCOCOOCCOOOCCOOOCOOCOOOOCOOOOOOCOOOCOCCCCCCCCOOCOCOCOCOOOCOOCCOCCOCCOOCCOCOOCOC
</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
5
fizz
7
8
fizz
10
11
fizz
13
busao
fizz
16
17
fizz
19
20
fizz
busao
23
fizz
25
26
fizz
busao
29
fizz
31
32
fizz
34
busao
fizz
37
38
fizz
40
41
fizzbusao
43
44
fizz
46
47
fizz
49
50
fizz
busao
53
fizz
55
56
fizz
58
59
fizz
61
62
fizz
busao
65
fizz
67
68
fizz
70
busao
fizz
73
74
fizz
76
77
fizz
79
80
fizzbusao
82
83
fizz
85
86
fizz
88
89
fizz
busao
92
fizz
94
95
fizz
97
98
fizz
100
101
fizzbusao
103
104
fizz
106
107
fizz
109
110
fizz
112
113
fizzbusao
115
116
fizz
118
119
fizzbusao
121
122
fizz
124
125
fizz
127
128
fizzbusao
130
131
fizz
133
134
fizz
136
137
fizz
busao
140
fizz
142
143
fizz
145
146
fizz
148
busao
fizz
151
152
fizz
154
155
fizz
157
158
fizz
160
161
fizz
busao
164
fizz
166
167
fizz
169
170
fizz
172
busao
fizz
175
176
fizz
178
179
fizz
181
busao
fizz
184
185
fizz
187
188
fizz
190
191
fizz
193
194
fizzbusao
196
197
fizz
199
200
fizz
202
203
fizz
busao
206
fizz
208
209
fizz
211
212
fizz
214
215
fizz
217
218
fizzbusao
220
221
fizz
223
224
fizz
busao
227
fizz
229
230
fizz
232
233
fizzbusao
235
236
fizz
238
239
fizz
241
242
fizz
busao
245
fizz
247
248
fizz
250
251
fizz
253
254
fizzbusao
256
257
fizz
259
260
fizz
262
263
fizzbusao
265
266
fizz
268
269
fizz
busao
272
fizz
274
275
fizz
277
278
fizz
busao
281
fizz
283
284
fizz
busao
287
fizz
289
290
fizz
292
293
fizz
295
296
fizzbusao
298
299
fizz
301
302
fizz
busao
305
fizz
307
308
fizz
310
311
fizz
313
314
fizz
busao
317
fizz
319
320
fizz
busao
323
fizz
325
326
fizz
328
329
fizz
331
332
fizzbusao
334
335
fizz
337
338
fizz
busao
341
fizz
343
344
fizz
346
347
fizz
349
350
fizzbusao
352
353
fizz
355
356
fizz
358
359
fizzbusao
361
362
fizz
364
365
fizz
367
368
fizz
370
371
fizz
373
busao
fizz
376
377
fizz
379
380
fizz
382
383
fizz
385
386
fizzbusao
388
389
fizz
391
392
fizz
394
busao
fizz
397
398
fizz
400
401
fizz
busao
404
fizz
406
407
fizz
409
410
fizz
412
413
fizzbusao
415
416
fizz
418
419
fizz
421
422
fizz
busao
425
fizz
427
428
fizz
430
431
fizz
433
434
fizz
436
busao
fizz
439
440
fizz
442
443
fizz
445
446
fizz
448
449
fizz
451
busao
fizz
454
455
fizz
457
458
fizz
busao
461
fizz
463
464
fizz
466
467
fizz
469
470
fizzbusao
472
473
fizz
475
476
fizz
478
479
fizz
busao
482
fizz
484
485
fizz
487
488
fizz
490
busao
fizz
493
494
fizz
496
497
fizz
499
500
fizz
busao
503
fizz
505
506
fizz
508
509
fizz
511
busao
fizz
514
515
fizz
517
518
fizz
busao
521
fizz
523
524
fizz
busao
527
fizz
529
530
fizz
532
533
fizz
535
busao
fizz
538
539
fizz
541
542
fizz
544
busao
fizz
547
548
fizz
550
551
fizz
553
554
fizz
556
busao
fizz
559
560
fizz
562
busao
fizz
565
566
fizz
568
569
fizz
571
572
fizz
busao
575
fizz
577
578
fizz
580
581
fizzbusao
583
584
fizz
586
587
fizzbusao
589
590
fizz
592
593
fizz
595
596
fizz
598
599
fizzbusao
601
602
fizz
604
605
fizz
607
608
fizzbusao
610
611
fizz
613
614
fizz
busao
617
fizz
619
620
fizz
busao
623
fizz
625
626
fizz
busao
629
fizz
631
632
fizz
634
635
fizz
637
638
fizz
640
641
fizzbusao
643
644
fizz
646
647
fizz
649
650
fizzbusao
652
653
fizz
655
656
fizz
658
659
fizz
busao
662
fizz
664
665
fizz
667
668
fizz
670
busao
fizz
673
674
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>COCCCOCCOCOCCCOOOOOOCCCCOOOOCCOCCCCCCOCCCOCOOCCOOOOCCCOCCCOOCOOOOOOCCOOCOCOCOCOOOCOCOOOOCCCCOOOCOOCCOCCCCOOCCOCCCOOCCCCOOCOCCOOOOOCCCCOCCCOOOCCOOOCCCOCCCOCOCCOOCOOOCCCOCOOOCOOCOCOCOOCCCOOCCCOOOCOCOCCOOOOOCOOOCOCCOCCCCOOOOOOOOCCCOOCOOCCOOCCCOOCCOOOOCCCOOOCOCCCCCCOCCOCOOCOOCOOOCCOCOCOOOOOCCCOOOOOCCCCCOCCCOCCCCCOOOOOCCOCOCOCOOOOOOOOCOOOCOCOCOCOOCOCOOCOOCCCCCCCCCCCCOCCOCCCC
</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
5
fizz
7
8
fizz
10
11
fizz
13
14
fizzbusao
16
17
fizz
19
busao
fizz
22
23
fizz
25
26
fizz
28
29
fizz
busao
32
fizz
34
35
fizz
37
38
fizz
40
41
fizz
43
44
fizz
46
47
fizzbusao
49
50
fizz
52
53
fizz
55
56
fizz
58
busao
fizz
61
62
fizz
64
busao
fizz
67
68
fizz
70
71
fizz
busao
74
fizz
76
77
fizz
79
80
fizzbusao
82
83
fizz
85
86
fizz
busao
89
fizz
91
92
fizz
94
95
fizz
97
busao
fizz
100
101
fizz
103
104
fizz
106
107
fizz
109
110
fizz
112
113
fizzbusao
115
116
fizz
118
119
fizz
121
122
fizz
124
125
fizzbusao
127
128
fizz
130
131
fizz
133
134
fizzbusao
136
137
fizz
139
140
fizz
142
143
fizz
busao
146
fizz
148
149
fizz
151
152
fizz
154
155
fizz
157
158
fizzbusao
160
161
fizz
163
164
fizz
166
167
fizzbusao
169
170
fizz
172
173
fizz
busao
176
fizz
178
179
fizz
181
182
fizz
184
185
fizzbusao
187
188
fizz
190
191
fizz
193
194
fizzbusao
196
197
fizz
199
200
fizz
202
busao
fizz
205
206
fizz
208
209
fizzbusao
211
212
fizz
214
215
fizz
217
218
fizz
220
busao
fizz
223
224
fizz
226
227
fizz
busao
230
fizz
232
233
fizz
235
236
fizzbusao
238
239
fizz
241
242
fizz
244
245
fizz
busao
248
fizz
250
251
fizz
253
254
fizz
busao
257
fizz
259
260
fizz
262
263
fizz
265
266
fizz
268
269
fizz
busao
272
fizz
274
275
fizz
277
278
fizzbusao
280
281
fizz
283
284
fizz
busao
287
fizz
289
290
fizz
292
293
fizzbusao
295
296
fizz
298
299
fizz
301
302
fizz
304
305
fizz
307
308
fizz
310
311
fizzbusao
313
314
fizz
316
317
fizz
319
busao
fizz
322
323
fizz
325
326
fizzbusao
328
329
fizz
331
332
fizzbusao
334
335
fizz
337
338
fizz
340
busao
fizz
343
344
fizz
346
347
fizz
busao
350
fizz
352
353
fizz
355
356
fizz
358
359
fizz
361
362
fizz
364
365
fizz
367
368
fizz
370
371
fizz
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>OCCCOCCCCOCCCOCCCOCCOOOCOCOCOCCCOOOOOCCOOOCCCCOOOOOCCCCOOCCCCCOOOOCOCOCOOOOCOCCCCOOOOCOCCOOCOCCOOCCOCOOCCCOOOCCCCCOOCOCCOOOOOOCOOCOOCCOCCOOOCCCCOCCOOCOOOCOCOOCCCOOCOOCCOOOOOOOCOCOCCCOOOOCOOCCCOCOCCCCOCOOCCOCCOOCCCCOOOCOCOCOCCOOOCOCOCCCOOCCOOOOOOOCCCCOOOCOOCCOCCOOCCCCCOCOCOCOOCCCOCOCOOCCOOCOCOCOCCCCOCOOOOOOCOCOCCOCOOCCOOCOCOOCCCOCOOOOOCOOOCCCCOOOCCOCOOCCOCOCCOOCOCCOOCOCCOCCOOOCCCCCCCCCCOCOOOOOOCOOOCCOOOCCCCOCOOOCOOCCCCCOOCOOOOOCOCCCOOOOCCOOCOOOCOOOCCCCOOCOOOOCOOCOOOCCCOOOOOCCCOOCOOCCOOOCOOOOOCCOCOOCCCOOCOOCOCOCOCCCCOCCOOCCCOCOOCOCCOCCCCCCCCOOOCOOOO
</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
5
fizz
7
8
fizz
10
11
fizz
13
14
fizz
16
17
fizzbusao
19
20
fizz
22
23
fizz
25
26
fizzbusao
28
29
fizz
31
32
fizz
34
35
fizzbusao
37
38
fizz
40
41
fizz
43
44
fizz
46
busao
fizz
49
50
fizz
52
53
fizz
55
busao
fizz
58
59
fizz
61
62
fizz
64
65
fizzbusao
67
68
fizz
70
71
fizz
73
busao
fizz
76
77
fizz
79
80
fizz
82
83
fizzbusao
85
86
fizz
88
89
fizz
91
92
fizzbusao
94
95
fizz
97
98
fizz
100
101
fizz
busao
104
fizz
106
107
fizz
109
110
fizz
112
113
fizz
115
busao
fizz
118
119
fizz
121
122
fizz
busao
125
fizz
127
128
fizz
130
busao
fizz
133
134
fizz
136
137
fizz
139
busao
fizz
142
143
fizz
145
146
fizz
148
149
fizz
151
busao
fizz
154
155
fizz
157
158
fizz
160
161
fizzbusao
163
164
fizz
166
167
fizz
169
busao
fizz
172
173
fizz
busao
176
fizz
178
179
fizz
181
182
fizz
184
busao
fizz
187
188
fizz
190
191
fizz
193
194
fizzbusao
196
197
fizz
199
200
fizz
202
203
fizz
205
206
fizz
208
busao
fizz
211
212
fizz
214
215
fizz
217
218
fizzbusao
220
221
fizz
223
224
fizz
226
227
fizzbusao
229
230
fizz
232
233
fizz
235
236
fizz
238
239
fizzbusao
241
242
fizz
244
busao
fizz
247
248
fizz
250
251
fizz
253
254
fizzbusao
256
257
fizz
259
260
fizz
262
263
fizz
265
266
fizz
268
busao
fizz
271
272
fizz
274
275
fizz
277
278
fizz
busao
281
fizz
283
284
fizz
286
287
fizz
busao
290
fizz
292
293
fizz
295
296
fizz
298
299
fizz
301
busao
fizz
304
305
fizz
busao
308
fizz
310
311
fizz
313
314
fizz
316
busao
fizz
319
320
fizz
322
323
fizz
325
busao
fizz
328
329
fizz
331
332
fizz
334
busao
fizz
337
338
fizz
340
341
fizz
343
344
fizzbusao
346
347
fizz
349
350
fizz
352
busao
fizz
355
356
fizz
358
359
fizz
361
362
fizz
busao
365
fizz
367
368
fizz
370
371
fizz
373
374
fizz
busao
377
fizz
379
380
fizz
382
383
fizz
385
386
fizz
388
389
fizz
391
busao
fizz
394
395
fizz
397
busao
fizz
400
401
fizz
403
404
fizzbusao
406
407
fizz
409
410
fizz
412
413
fizz
415
busao
fizz
418
419
fizz
421
422
fizz
424
425
fizz
busao
428
fizz
430
431
fizz
433
434
fizz
busao
437
fizz
439
440
fizz
442
busao
fizz
445
446
fizz
448
449
fizzbusao
451
452
fizz
454
455
fizz
457
458
fizz
busao
461
fizz
463
464
fizz
466
busao
fizz
469
470
fizz
472
473
fizz
busao
476
fizz
478
479
fizz
481
482
fizz
busao
485
fizz
487
488
fizz
490
491
fizzbusao
493
494
fizz
496
497
fizz
busao
500
fizz
502
503
fizz
505
506
fizz
508
busao
fizz
511
512
fizz
514
515
fizz
517
518
fizz
520
busao
fizz
523
524
fizz
526
527
fizz
529
530
fizz
busao
533
fizz
535
536
fizz
538
539
fizz
541
542
fizz
544
545
fizz
547
busao
fizz
550
551
fizz
553
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>COCOOCCCCCOOCCOOCCOCOCCCCCOOCCOCOCOOOOCOOCOCCCOCCCOCOCCCCCOCOOOCOOCCOOOCOOOCCOOOOCOOCCOCCCOCOOCCCOCOCOOCOOCCOCOCCCCCCOCOOCOCCOOOOCOOCOCOCCOOCCCOCCCCCOCOOCCOOCCCOOOOCOCOOOOCOOOCOCCOCCCOCOCCOCOOOCOCCOOCOOOCCOCCCOOCCOOCOCCOCCOOOOOCOCOOCOOCOOOCOCCCCOOOCCOOCCCOOOOCOCOCCOOOOCOOOOCOCOCOOOOCCOOOCCOCCCOCCOOCOOCCOOOCOOOCCCOOOOCCOCCCCCCOOCOCCCOOCOOOOOCOCCCCOOCOOOCOOCOOOCOCOCCOCCCCCOCOCOCCCOCCCCOOOOCOOOCOOCOOCCOCOOOCCOOOCOCOOOCOOCOCCCOOCCOCOCOCOOOOOOOOCCOOOOOCCOOOOCOCOCOOOCCOCOCCCCCCCOOCOOOOCOCOCCOCOOCCOCOOCOOOOCOCCCOCOCOOOOCCCOCOOCCOCOCCCCOCCCOCOOOOOCOCOOOOOOOCOOOCOCOOCOOOOCOCOOCOCOOOCCOCCCOCCCOCOOOCOOOOCOOOOCOCCOOOOOCCCCOOCOCCCOOCOOCCCCCOOOCOOOCCOCCOOCCCCOCOOOOCCCCCOOCCOOCCCOCOOOCCOCOOOOCCOCCOCCCCOOOOCOCOOOOCOCOOOCOCOCCCOCOOCCCOOCOOOOCCCCOCCCOCOOOCOOCOCCCCCCCCCCOO
</text>
      </stdin>
      <expected>
                <text>1
2
fizz
4
5
fizz
7
8
fizz
10
11
fizzbusao
13
14
fizz
16
17
fizz
19
20
fizz
22
23
fizz
25
26
fizzbusao
28
29
fizz
31
32
fizz
34
35
fizzbusao
37
38
fizz
40
41
fizz
busao
44
fizz
46
47
fizz
49
50
fizz
52
53
fizz
55
56
fizz
58
59
fizz
busao
62
fizz
64
65
fizz
67
68
fizzbusao
70
71
fizz
73
74
fizzbusao
76
77
fizz
79
80
fizz
82
busao
fizz
85
86
fizz
88
89
fizz
91
92
fizz
busao
95
fizz
97
98
fizz
100
101
fizz
103
104
fizzbusao
106
107
fizz
109
110
fizz
112
113
fizz
115
116
fizz
118
119
fizzbusao
121
122
fizz
124
125
fizz
127
busao
fizz
130
131
fizz
133
134
fizz
busao
137
fizz
139
140
fizz
142
143
fizz
145
146
fizz
148
149
fizz
151
busao
fizz
154
155
fizz
157
158
fizz
160
161
fizzbusao
163
164
fizz
166
167
fizz
busao
170
fizz
172
173
fizz
busao
176
fizz
178
179
fizz
181
182
fizz
184
185
fizz
187
188
fizzbusao
190
191
fizz
193
194
fizz
196
197
fizzbusao
199
200
fizz
202
203
fizz
205
busao
fizz
208
209
fizz
211
212
fizz
214
215
fizz
busao
218
fizz
220
221
fizz
223
224
fizz
busao
227
fizz
229
230
fizz
232
233
fizzbusao
235
236
fizz
238
239
fizz
busao
242
fizz
244
245
fizz
247
248
fizz
250
251
fizzbusao
253
254
fizz
256
257
fizz
259
260
fizzbusao
262
263
fizz
265
266
fizz
268
busao
fizz
271
272
fizz
274
275
fizzbusao
277
278
fizz
280
281
fizz
busao
284
fizz
286
287
fizz
289
290
fizz
292
293
fizz
busao
296
fizz
298
299
fizz
301
302
fizz
304
busao
fizz
307
308
fizz
310
busao
fizz
313
314
fizz
316
317
fizz
319
320
fizzbusao
322
323
fizz
325
326
fizz
328
329
fizz
331
332
fizz
334
335
fizzbusao
337
338
fizz
340
341
fizzbusao
343
344
fizz
346
347
fizz
349
350
fizz
352
busao
fizz
355
356
fizz
358
359
fizzbusao
361
362
fizz
364
365
fizz
367
368
fizz
370
371
fizz
373
busao
fizz
376
377
fizz
379
380
fizz
382
383
fizz
385
386
fizz
busao
389
fizz
391
392
fizz
busao
395
fizz
397
398
fizz
400
401
fizz
busao
404
fizz
406
407
fizz
409
410
fizzbusao
412
413
fizz
415
416
fizz
busao
419
fizz
421
422
fizz
424
425
fizz
427
busao
fizz
430
431
fizz
433
434
fizz
436
437
fizzbusao
439
440
fizz
442
busao
fizz
445
446
fizz
448
449
fizzbusao
451
452
fizz
454
455
fizz
busao
458
fizz
460
461
fizz
463
464
fizzbusao
466
467
fizz
469
470
fizz
472
473
fizz
475
476
fizz
478
479
fizz
busao
482
fizz
484
485
fizz
487
busao
fizz
490
491
fizz
493
494
fizz
496
497
fizz
busao
500
fizz
502
503
fizz
busao
506
fizz
508
509
fizz
511
512
fizz
514
515
fizzbusao
517
518
fizz
520
521
fizz
523
524
fizzbusao
526
527
fizz
529
530
fizz
532
533
fizz
535
536
fizz
538
539
fizz
busao
542
fizz
544
545
fizz
busao
548
fizz
550
551
fizz
busao
554
fizz
556
557
fizz
busao
560
fizz
562
563
fizz
565
566
fizzbusao
568
569
fizz
571
572
fizz
busao
575
fizz
577
578
fizz
580
581
fizz
busao
584
fizz
586
587
fizz
589
590
fizz
592
593
fizz
busao
596
fizz
598
599
fizz
601
busao
fizz
604
605
fizz
607
608
fizz
busao
611
fizz
613
614
fizz
616
617
fizz
busao
620
fizz
622
623
fizz
625
626
fizz
628
busao
fizz
631
632
fizz
634
635
fizz
637
638
fizz
busao
641
fizz
643
644
fizz
646
647
fizz
busao
650
fizz
652
653
fizz
655
656
fizz
658
busao
fizz
661
662
fizz
664
665
fizz
667
668
fizz
670
671
fizz
673
busao
fizz
676
677
fizz
679
680
fizz
682
busao
fizz
685
686
fizz
688
689
fizz
691
busao
fizz
694
695
fizz
697
698
fizz
700
701
fizzbusao
703
704
fizz
706
707
fizz
busao
710
fizz
712
713
fizz
715
716
fizzbusao
718
719
fizz
721
722
fizz
724
725
fizz
727
728
fizzbusao
730
731
fizz
733
734
fizz
736
737
fizz
busao
740
fizz
742
743
fizz
745
746
fizz
748
busao
fizz
751
752
fizz
754
755
fizz
757
758
fizz
760
761
fizz
763
764
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>C
</text>
      </stdin>
      <expected>
                <text>1
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>O
</text>
      </stdin>
      <expected>
                <text>1
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
    <tags>
      <tag><text>adhoc</text>
</tag>
      <tag><text>remainder</text>
</tag>
    </tags>
  </question>

</quiz>