<quiz>
  <question type="coderunner">
    <name>
      <text>Infixa para p&#243;s-fixa
</text>
    </name>
    <questiontext format="html">
      <text>
<![CDATA[
<p>
<b></b><br />
Dona Ros&#226;ngela est&#225; escrevendo um livro matem&#225;tico contendo apenas em express&#245;es matem&#225;ticas em nota&#231;&#227;o p&#243;s-fixa. Sabemos que as express&#245;es aritm&#233;ticas que estudamos est&#227;o em nota&#231;&#227;o infixa, isto &#233;, em opera&#231;&#245;es matem&#225;ticas bin&#225;rias, o operador \( y \) &#233; apresentado entre dois operandos (\( o_1 \) e \( o_2 \)), sob a forma \( o_1 y o_2 \).

Em seus estudos, Dona Ros&#226;ngela considera as seguintes opera&#231;&#245;es matem&#225;ticas bin&#225;rias:

<ul>
<li> ^ (pot&#234;ncia);</li>
<li> / (divis&#227;o);</li>
<li> * (multiplica&#231;&#227;o);</li>
<li> - (subtra&#231;&#227;o);</li>
<li> + (soma).</li>
</ul>

Al&#233;m disso, os operandos podem ser representados por vari&#225;veis min&#250;sculas (de \( a \) at&#233; \( z \)) ou constantes de um &#250;nico d&#237;gito num&#233;rico. Nessas express&#245;es, podem existir par&#234;nteses para preservar a prioridade de determinadas opera&#231;&#245;es, uma vez que alguns operadores possuem prioridade na aus&#234;ncia de par&#234;nteses (por exemplo, multiplica&#231;&#227;o tem prioridade em rela&#231;&#227;o &#224; soma e subtra&#231;&#227;o).

Por exemplo, a express&#227;o em nota&#231;&#227;o p&#243;s-fixa ab*cd+2^/ &#233; obtida a partir da convers&#227;o da express&#227;o em nota&#231;&#227;o infixa a*4/(c+d)^2. Ajude Dona Ros&#226;ngela nessa tarefa e elabore um programa que converta uma express&#227;o em nota&#231;&#227;o infixa para a nota&#231;&#227;o p&#243;s-fixa.

</p>

<p>
<b>Entrada</b><br />
A entrada consiste de uma express&#227;o matem&#225;tica \( 1 \leq | s | \leq 10^3 \) em nota&#231;&#227;o infixa.

</p>

<p>
<b>Saida</b><br />
Imprima a express&#227;o correspondente em nota&#231;&#227;o p&#243;s-fixa.

</p>
]]>
</text>
    </questiontext>
    <generalfeedback format="html">
      <text />
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber />
    <coderunnertype>cpp_program</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0, 0, 10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload />
    <useace />
    <resultcolumns />
    <template />
    <iscombinatortemplate />
    <allowmultiplestdins />
    <answer>
<![CDATA[#include<bits/stdc++.h>

using namespace std;

map<char,int> op;

void prioridade()
{
	op['+']=1;
	op['-']=1;
	op['/']=2;
	op['*']=2;
	op['^']=3;
}

void convertInfixa2PosFixa(char exp[])
{
    stack<char> pilha;

    for(int i=0; exp[i] != '\0';i++)
    {
        if((exp[i] >= 'a' && exp[i] <= 'z') || (exp[i] >= '0' && exp[i] <= '9'))
        {
            printf("%c",exp[i]);
        }
        else
        {
            if(exp[i] == '+' || exp[i]== '-' || exp[i]== '*' || exp[i]== '/' || exp[i]== '^')
            {
                if(!pilha.empty())
                {
                    while(op[exp[i]] <= op[pilha.top()])
                    {
                        printf("%c",(char) pilha.top());
                        
                        pilha.pop();
                        
                        if(pilha.empty())
                            break;
                    }
                }
		pilha.push(exp[i]);
            }
            else
            {
                if(exp[i]=='(')
                {
                    pilha.push('(');
                }
                else
                {
                    if(exp[i]==')')
                    {
                        while(pilha.top()!='(')
                        {
                            printf("%c",(char) pilha.top());
                            pilha.pop();
                        }
                        pilha.pop();
                    }
                }
            }
        }
    }
    
    while(!pilha.empty())
    {
        printf("%c",(char) pilha.top());
        pilha.pop();
    }

    printf("\n");
}

int main()
{
    
    prioridade();

    char exp[101];
    
    scanf("%s",exp);

    convertInfixa2PosFixa(exp);

    return 0;
}
]]>
</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre />
    <language />
    <acelang />
    <sandbox />
    <grader />
    <cputimelimitsecs />
    <memlimitmb />
    <sandboxparams />
    <templateparams />
    <hoisttemplateparams>1</hoisttemplateparams>
    <twigall>0</twigall>
    <uiplugin />
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex />
    <filenamesexplain />
    <displayfeedback>1</displayfeedback>
    <testcases>
    <testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>a+b
</text>
      </stdin>
      <expected>
                <text>ab+
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>a*2+b
</text>
      </stdin>
      <expected>
                <text>a2*b+
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>b/(c+1)
</text>
      </stdin>
      <expected>
                <text>bc1+/
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>(a+b)/(c-3)
</text>
      </stdin>
      <expected>
                <text>ab+c3-/
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>(a*2+c-d)/2
</text>
      </stdin>
      <expected>
                <text>a2*c+d-2/
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>(2*4/a^b)/(2*c)
</text>
      </stdin>
      <expected>
                <text>24*ab^/2c*/
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>a*b-c*d^e/f+g*h
</text>
      </stdin>
      <expected>
                <text>ab*cde^*f/-gh*+
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>a*(b+c)*(d-g)*h
</text>
      </stdin>
      <expected>
                <text>abc+*dg-*h*
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>a+b*c^d-e
</text>
      </stdin>
      <expected>
                <text>abcd^*+e-
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>(a+b)*c/(2-r)/2
</text>
      </stdin>
      <expected>
                <text>ab+c*2r-/2/
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>(a*(b+3*(w-7/(a+b))))
</text>
      </stdin>
      <expected>
                <text>ab3w7ab+/-*+*
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>(((t^2*a/b)-e)*2)+3-r
</text>
      </stdin>
      <expected>
                <text>t2^a*b/e-2*3+r-
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>t^3^5^7^c+p
</text>
      </stdin>
      <expected>
                <text>t3^5^7^c^p+
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>a*(2+3)
</text>
      </stdin>
      <expected>
                <text>a23+*
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>b*(2/3+1)
</text>
      </stdin>
      <expected>
                <text>b23/1+*
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>(2*(4/a)^b*1+2*(3+5))/(2*(c-3))
</text>
      </stdin>
      <expected>
                <text>24a/b^*1*235+*+2c3-*/
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>((b+2)^3)/4^4*(3+1)
</text>
      </stdin>
      <expected>
                <text>b2+3^44^/31+*
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>(a)
</text>
      </stdin>
      <expected>
                <text>a
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>(a+1)
</text>
      </stdin>
      <expected>
                <text>a1+
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>b*(1)
</text>
      </stdin>
      <expected>
                <text>b1*
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>6+4/b^3
</text>
      </stdin>
      <expected>
                <text>64b3^/+
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>a+b*a/d-u/b
</text>
      </stdin>
      <expected>
                <text>aba*d/+ub/-
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>(a+2)^5-(b*c)/6+(a+2)^5-(b*c)/6-(a+2)^5-(b*c)/6+(a+2)^5-(b*c)/6+(2/3+x+1)
</text>
      </stdin>
      <expected>
                <text>a2+5^bc*6/-a2+5^+bc*6/-a2+5^-bc*6/-a2+5^+bc*6/-23/x+1++
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>a-b/c^d-e-f-g^h*i/j*k^l/m*n^o*p+q*r+s+t+u^v*w/x+y+z
</text>
      </stdin>
      <expected>
                <text>abcd^/-e-f-gh^i*j/kl^*m/no^*p*-qr*+s+t+uv^w*x/+y+z+
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>a-b^c/d+e-f^g^h-i+j+k+l*m*n^o*p^q/r+s/t^u^v/w-x-y/z
</text>
      </stdin>
      <expected>
                <text>abc^d/-e+fg^h^-i-j+k+lm*no^*pq^*r/+stu^v^/w/+x-yz/-
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>a/b-c*d/e*f/g+h+i-j+k/l/m-n^o^p^q^r+s+t/u*v+w/x*y+z
</text>
      </stdin>
      <expected>
                <text>ab/cd*e/f*g/-h+i+j-kl/m/+no^p^q^r^-s+tu/v*+wx/y*+z+
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>((a+b)*4-(1-e))^(f-g)
</text>
      </stdin>
      <expected>
                <text>ab+4*1e--fg-^
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>((a+t)*((b+(a+c))^(c+d)))
</text>
      </stdin>
      <expected>
                <text>at+bac++cd+^*
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>a^b+c^d+e*f-g^h-i+j-k/l*m^n+o+p+q+r+s/t/u/v-w*x^y/z
</text>
      </stdin>
      <expected>
                <text>ab^cd^+ef*+gh^-i-j+kl/mn^*-o+p+q+r+st/u/v/+wxy^*z/-
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase><testcase hiderestiffail="0" mark="1.0000000" testtype="0" useasexample="0">
      <testcode>
                <text />
      </testcode>
      <stdin>
                <text>(a*b+1)/(d+c*7)^(a-2+1+t*(a+b))+6-i^2^(t+3)-a*(c+d/2)
</text>
      </stdin>
      <expected>
                <text>ab*1+dc7*+a2-1+tab+*+^/6+i2^t3+^-acd2/+*-
</text>
      </expected>
      <extra>
                <text />
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase></testcases>
  </question>

</quiz>